%--------------------------------------------------------------------------
% Proof Results of "Torque_phi.tptp"
%--------------------------------------------------------------------------

Z3 version 4.4.2 - 64 bit
---------------------------------------------------------------------------
SZS status Theorem for /home/omar/Desktop/Omar - ex/My Problems/Helicopter/Torque_phi.tptp

SZS output start CNFRefutation for /home/omar/Desktop/Omar - ex/My Problems/Helicopter/Torque_phi.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_left_divide_mul_neg, axiom, (Y * Z < X | Y <= X / Z | 0 <= Z)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880,
            Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

fof('Torque_phi', conjecture,
    (! [T_phi, TM, TR, A, B, Ly_M, Lz_M, Lz_R] :
       ((-7/200 <= T_phi & T_phi <= 367/2500 & TM = 1357143/10000 &
         TR = 12/5 & A = 1047/5000 & B = 2443/10000 & Ly_M = 0 &
         Lz_M = -1/4 & Lz_R = 0) =>
        abs(T_phi) <=
        abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
        (-Ly_M * cos(A) * cos(B) - Lz_M * cos(A) * sin(B)) - Lz_R * TR))).

fof(subgoal_0, plain,
    (! [T_phi, TM, TR, A, B, Ly_M, Lz_M, Lz_R] :
       ((-7/200 <= T_phi & T_phi <= 367/2500 & TM = 1357143/10000 &
         TR = 12/5 & A = 1047/5000 & B = 2443/10000 & Ly_M = 0 &
         Lz_M = -1/4 & Lz_R = 0) =>
        abs(T_phi) <=
        abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
        (-Ly_M * cos(A) * cos(B) - Lz_M * cos(A) * sin(B)) - Lz_R * TR)),
    inference(strip, [], ['Torque_phi'])).

fof(negate_0_0, plain,
    (~ ! [T_phi, TM, TR, A, B, Ly_M, Lz_M, Lz_R] :
         ((-7/200 <= T_phi & T_phi <= 367/2500 & TM = 1357143/10000 &
           TR = 12/5 & A = 1047/5000 & B = 2443/10000 & Ly_M = 0 &
           Lz_M = -1/4 & Lz_R = 0) =>
          abs(T_phi) <=
          abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
          (-Ly_M * cos(A) * cos(B) - Lz_M * cos(A) * sin(B)) - Lz_R * TR)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [A, B, Ly_M, Lz_M, Lz_R, TM, TR, T_phi] :
       (abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
        (-Ly_M * cos(A) * cos(B) - Lz_M * cos(A) * sin(B)) - Lz_R * TR <
        abs(T_phi) & T_phi <= 367/2500 & -7/200 <= T_phi & A = 1047/5000 &
        B = 2443/10000 & Ly_M = 0 & Lz_M = -1/4 & Lz_R = 0 &
        TM = 1357143/10000 & TR = 12/5)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 <
     abs(skoT_phiC1) & -7/200 <= skoT_phiC1 & skoT_phiC1 <= 367/2500 &
     skoAC1 = 1047/5000 & skoBC1 = 2443/10000 & skoLy_MC1 = 0 &
     skoLz_MC1 = -1/4 & skoLz_RC1 = 0 & skoTMC1 = 1357143/10000 &
     skoTRC1 = 12/5), inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 <
     abs(skoT_phiC1)), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoTRC1 = 12/5),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoLz_RC1 = 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (skoBC1 = 2443/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (skoAC1 = 1047/5000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_7, plain, (skoLz_MC1 = -1/4),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_8, plain, (skoLy_MC1 = 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_9, plain, (skoTMC1 = 1357143/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_10, plain, (-7/200 <= skoT_phiC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_11, plain, (skoT_phiC1 <= 367/2500),
    inference(conjunct, [], [normalize_0_1])).

cnf(refute_0_0, plain,
    (~ lgen(0, X_000029, cos(X_000028)) | X_000029 <= cos(X_000028)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_1, plain,
    (~ lgen(0, X_000029,
            1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24 - X_000028 ^ 6 / 720 +
            X_000028 ^ 8 / 40320 - X_000028 ^ 10 / 3628800) |
     lgen(0, X_000029, cos(X_000028))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_2, plain,
    (~ lgen(0, X_000029,
            1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24 - X_000028 ^ 6 / 720 +
            X_000028 ^ 8 / 40320 - X_000028 ^ 10 / 3628800) |
     X_000029 <= cos(X_000028)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (1 +
     X_000028 *
     (X_000028 *
      (-1/2 +
       X_000028 *
       (X_000028 *
        (1/24 +
         X_000028 *
         (X_000028 *
          (-1/720 +
           X_000028 *
           (X_000028 *
            (1/40320 + X_000028 * (X_000028 * -1/3628800))))))))) <
     X_000029 | X_000029 <= cos(X_000028)),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (1 +
     1047/5000 *
     (1047/5000 *
      (-1/2 +
       1047/5000 *
       (1047/5000 *
        (1/24 +
         1047/5000 *
         (1047/5000 *
          (-1/720 +
           1047/5000 *
           (1047/5000 *
            (1/40320 + 1047/5000 * (1047/5000 * -1/3628800))))))))) <
     skoT_phiC1 *
     -852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     skoT_phiC1 *
     -852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     cos(1047/5000)), inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_6, plain,
    (~ lgen(0, X_000168, sin(X_000167)) | X_000168 <= sin(X_000167)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_7, plain,
    (X_000167 < 0 |
     ~ lgen(0, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800) | lgen(0, X_000168, sin(X_000167))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_8, plain,
    (X_000167 < 0 |
     ~ lgen(0, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800) | X_000168 <= sin(X_000167)),
    inference(resolve, [], [refute_0_7, refute_0_6])).

cnf(refute_0_9, plain,
    (X_000167 *
     (1 +
      X_000167 *
      (X_000167 *
       (-1/6 +
        X_000167 *
        (X_000167 *
         (1/120 +
          X_000167 *
          (X_000167 *
           (-1/5040 +
            X_000167 *
            (X_000167 *
             (1/362880 + X_000167 * (X_000167 * -1/39916800)))))))))) <
     X_000168 | X_000167 < 0 | X_000168 <= sin(X_000167)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_11, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_12, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     sin(1047/5000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_13, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_14, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_15, plain,
    (skoT_phiC1 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     0 <=
     cos(1047/5000) * -1357143/40000 +
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_16, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_17, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (skoT_phiC1 * -1) | skoT_phiC1 * -1 <= 0 |
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) <=
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) /
     (skoT_phiC1 * -1)), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_18, plain,
    (~ lgen(0, sqrt(X_000086), X_000087) | sqrt(X_000086) <= X_000087),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_19, plain,
    (X_000086 < 0 | ~ lgen(0, (X_000086 + 1) / 2, X_000087) |
     lgen(0, sqrt(X_000086), X_000087)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_20, plain,
    (X_000086 < 0 | ~ lgen(0, (X_000086 + 1) / 2, X_000087) |
     sqrt(X_000086) <= X_000087),
    inference(resolve, [], [refute_0_19, refute_0_18])).

cnf(refute_0_21, plain,
    (X_000086 < 0 | X_000087 < 1/2 + X_000086 * 1/2 |
     sqrt(X_000086) <= X_000087),
    inference(arithmetic, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1) <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1)),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_23, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 <
     abs(skoT_phiC1)), inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_24, plain, (skoTRC1 = 12/5),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_25, plain, (0 * skoTRC1 = 0 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_26, plain,
    (0 * skoTRC1 != 0 * skoTRC1 | skoTRC1 != 12/5 |
     0 * skoTRC1 = 0 * 12/5), introduced(tautology, [equality])).

cnf(refute_0_27, plain, (skoTRC1 != 12/5 | 0 * skoTRC1 = 0 * 12/5),
    inference(resolve, [], [refute_0_25, refute_0_26])).

cnf(refute_0_28, plain, (0 * skoTRC1 = 0 * 12/5),
    inference(resolve, [], [refute_0_24, refute_0_27])).

cnf(refute_0_29, plain, (skoLz_RC1 = 0),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_30, plain, (skoLz_RC1 * skoTRC1 = skoLz_RC1 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_31, plain,
    (skoLz_RC1 * skoTRC1 != skoLz_RC1 * skoTRC1 | skoLz_RC1 != 0 |
     skoLz_RC1 * skoTRC1 = 0 * skoTRC1),
    introduced(tautology, [equality])).

cnf(refute_0_32, plain,
    (skoLz_RC1 != 0 | skoLz_RC1 * skoTRC1 = 0 * skoTRC1),
    inference(resolve, [], [refute_0_30, refute_0_31])).

cnf(refute_0_33, plain, (skoLz_RC1 * skoTRC1 = 0 * skoTRC1),
    inference(resolve, [], [refute_0_29, refute_0_32])).

cnf(refute_0_34, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_35, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_36, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_34, refute_0_35])).

cnf(refute_0_37, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_38, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_36, refute_0_37])).

cnf(refute_0_39, plain,
    (0 * skoTRC1 != 0 * 12/5 | skoLz_RC1 * skoTRC1 != 0 * skoTRC1 |
     skoLz_RC1 * skoTRC1 = 0 * 12/5), inference(subst, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (0 * skoTRC1 != 0 * 12/5 | skoLz_RC1 * skoTRC1 = 0 * 12/5),
    inference(resolve, [], [refute_0_33, refute_0_39])).

cnf(refute_0_41, plain, (skoLz_RC1 * skoTRC1 = 0 * 12/5),
    inference(resolve, [], [refute_0_28, refute_0_40])).

cnf(refute_0_42, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_43, plain,
    (skoLz_RC1 * skoTRC1 != 0 * 12/5 |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5),
    introduced(tautology, [equality])).

cnf(refute_0_44, plain,
    (skoLz_RC1 * skoTRC1 != 0 * 12/5 |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5),
    inference(resolve, [], [refute_0_41, refute_0_44])).

cnf(refute_0_46, plain, (skoBC1 = 2443/10000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_47, plain, (sin(skoBC1) = sin(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_48, plain,
    (sin(skoBC1) != sin(skoBC1) | skoBC1 != 2443/10000 |
     sin(skoBC1) = sin(2443/10000)), introduced(tautology, [equality])).

cnf(refute_0_49, plain,
    (skoBC1 != 2443/10000 | sin(skoBC1) = sin(2443/10000)),
    inference(resolve, [], [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain, (sin(skoBC1) = sin(2443/10000)),
    inference(resolve, [], [refute_0_46, refute_0_49])).

cnf(refute_0_51, plain,
    (-1/4 * cos(1047/5000) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(skoBC1)), introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (-1/4 * cos(1047/5000) * sin(skoBC1) !=
     -1/4 * cos(1047/5000) * sin(skoBC1) | sin(skoBC1) != sin(2443/10000) |
     -1/4 * cos(1047/5000) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (sin(skoBC1) != sin(2443/10000) |
     -1/4 * cos(1047/5000) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain,
    (-1/4 * cos(1047/5000) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain, (skoAC1 = 1047/5000),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_56, plain, (cos(skoAC1) = cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_57, plain,
    (cos(skoAC1) != cos(skoAC1) | skoAC1 != 1047/5000 |
     cos(skoAC1) = cos(1047/5000)), introduced(tautology, [equality])).

cnf(refute_0_58, plain,
    (skoAC1 != 1047/5000 | cos(skoAC1) = cos(1047/5000)),
    inference(resolve, [], [refute_0_56, refute_0_57])).

cnf(refute_0_59, plain, (cos(skoAC1) = cos(1047/5000)),
    inference(resolve, [], [refute_0_55, refute_0_58])).

cnf(refute_0_60, plain, (-1/4 * cos(skoAC1) = -1/4 * cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_61, plain,
    (-1/4 * cos(skoAC1) != -1/4 * cos(skoAC1) |
     cos(skoAC1) != cos(1047/5000) |
     -1/4 * cos(skoAC1) = -1/4 * cos(1047/5000)),
    introduced(tautology, [equality])).

cnf(refute_0_62, plain,
    (cos(skoAC1) != cos(1047/5000) |
     -1/4 * cos(skoAC1) = -1/4 * cos(1047/5000)),
    inference(resolve, [], [refute_0_60, refute_0_61])).

cnf(refute_0_63, plain, (-1/4 * cos(skoAC1) = -1/4 * cos(1047/5000)),
    inference(resolve, [], [refute_0_59, refute_0_62])).

cnf(refute_0_64, plain, (skoLz_MC1 = -1/4),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_65, plain,
    (skoLz_MC1 * cos(skoAC1) = skoLz_MC1 * cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_66, plain,
    (skoLz_MC1 * cos(skoAC1) != skoLz_MC1 * cos(skoAC1) |
     skoLz_MC1 != -1/4 | skoLz_MC1 * cos(skoAC1) = -1/4 * cos(skoAC1)),
    introduced(tautology, [equality])).

cnf(refute_0_67, plain,
    (skoLz_MC1 != -1/4 | skoLz_MC1 * cos(skoAC1) = -1/4 * cos(skoAC1)),
    inference(resolve, [], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain, (skoLz_MC1 * cos(skoAC1) = -1/4 * cos(skoAC1)),
    inference(resolve, [], [refute_0_64, refute_0_67])).

cnf(refute_0_69, plain,
    (-1/4 * cos(skoAC1) != -1/4 * cos(1047/5000) |
     skoLz_MC1 * cos(skoAC1) != -1/4 * cos(skoAC1) |
     skoLz_MC1 * cos(skoAC1) = -1/4 * cos(1047/5000)),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_70, plain,
    (-1/4 * cos(skoAC1) != -1/4 * cos(1047/5000) |
     skoLz_MC1 * cos(skoAC1) = -1/4 * cos(1047/5000)),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain, (skoLz_MC1 * cos(skoAC1) = -1/4 * cos(1047/5000)),
    inference(resolve, [], [refute_0_63, refute_0_70])).

cnf(refute_0_72, plain,
    (skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_73, plain,
    (skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) |
     skoLz_MC1 * cos(skoAC1) != -1/4 * cos(1047/5000) |
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(skoBC1)),
    introduced(tautology, [equality])).

cnf(refute_0_74, plain,
    (skoLz_MC1 * cos(skoAC1) != -1/4 * cos(1047/5000) |
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(skoBC1)),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain,
    (skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(skoBC1)),
    inference(resolve, [], [refute_0_71, refute_0_74])).

cnf(refute_0_76, plain,
    (-1/4 * cos(1047/5000) * sin(skoBC1) !=
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -1/4 * cos(1047/5000) * sin(skoBC1) |
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_77, plain,
    (-1/4 * cos(1047/5000) * sin(skoBC1) !=
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_75, refute_0_76])).

cnf(refute_0_78, plain,
    (skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_54, refute_0_77])).

cnf(refute_0_79, plain,
    (-0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_80, plain,
    (skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) |
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_81, plain,
    (skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_79, refute_0_80])).

cnf(refute_0_82, plain,
    (-0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_78, refute_0_81])).

cnf(refute_0_83, plain, (cos(skoBC1) = cos(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_84, plain,
    (cos(skoBC1) != cos(skoBC1) | skoBC1 != 2443/10000 |
     cos(skoBC1) = cos(2443/10000)), introduced(tautology, [equality])).

cnf(refute_0_85, plain,
    (skoBC1 != 2443/10000 | cos(skoBC1) = cos(2443/10000)),
    inference(resolve, [], [refute_0_83, refute_0_84])).

cnf(refute_0_86, plain, (cos(skoBC1) = cos(2443/10000)),
    inference(resolve, [], [refute_0_46, refute_0_85])).

cnf(refute_0_87, plain,
    (-0 * cos(1047/5000) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(skoBC1)), introduced(tautology, [refl])).

cnf(refute_0_88, plain,
    (-0 * cos(1047/5000) * cos(skoBC1) !=
     -0 * cos(1047/5000) * cos(skoBC1) | cos(skoBC1) != cos(2443/10000) |
     -0 * cos(1047/5000) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_89, plain,
    (cos(skoBC1) != cos(2443/10000) |
     -0 * cos(1047/5000) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_87, refute_0_88])).

cnf(refute_0_90, plain,
    (-0 * cos(1047/5000) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_86, refute_0_89])).

cnf(refute_0_91, plain, (-0 * cos(skoAC1) = -0 * cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_92, plain,
    (-0 * cos(skoAC1) != -0 * cos(skoAC1) | cos(skoAC1) != cos(1047/5000) |
     -0 * cos(skoAC1) = -0 * cos(1047/5000)),
    introduced(tautology, [equality])).

cnf(refute_0_93, plain,
    (cos(skoAC1) != cos(1047/5000) |
     -0 * cos(skoAC1) = -0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_91, refute_0_92])).

cnf(refute_0_94, plain, (-0 * cos(skoAC1) = -0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_59, refute_0_93])).

cnf(refute_0_95, plain, (skoLy_MC1 = 0),
    inference(canonicalize, [], [normalize_0_8])).

cnf(refute_0_96, plain, (-skoLy_MC1 = -skoLy_MC1),
    introduced(tautology, [refl])).

cnf(refute_0_97, plain,
    (-skoLy_MC1 != -skoLy_MC1 | skoLy_MC1 != 0 | -skoLy_MC1 = -0),
    introduced(tautology, [equality])).

cnf(refute_0_98, plain, (skoLy_MC1 != 0 | -skoLy_MC1 = -0),
    inference(resolve, [], [refute_0_96, refute_0_97])).

cnf(refute_0_99, plain, (-skoLy_MC1 = -0),
    inference(resolve, [], [refute_0_95, refute_0_98])).

cnf(refute_0_100, plain,
    (-skoLy_MC1 * cos(skoAC1) = -skoLy_MC1 * cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_101, plain,
    (-skoLy_MC1 * cos(skoAC1) != -skoLy_MC1 * cos(skoAC1) |
     -skoLy_MC1 != -0 | -skoLy_MC1 * cos(skoAC1) = -0 * cos(skoAC1)),
    introduced(tautology, [equality])).

cnf(refute_0_102, plain,
    (-skoLy_MC1 != -0 | -skoLy_MC1 * cos(skoAC1) = -0 * cos(skoAC1)),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain, (-skoLy_MC1 * cos(skoAC1) = -0 * cos(skoAC1)),
    inference(resolve, [], [refute_0_99, refute_0_102])).

cnf(refute_0_104, plain,
    (-0 * cos(skoAC1) != -0 * cos(1047/5000) |
     -skoLy_MC1 * cos(skoAC1) != -0 * cos(skoAC1) |
     -skoLy_MC1 * cos(skoAC1) = -0 * cos(1047/5000)),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_105, plain,
    (-0 * cos(skoAC1) != -0 * cos(1047/5000) |
     -skoLy_MC1 * cos(skoAC1) = -0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_103, refute_0_104])).

cnf(refute_0_106, plain, (-skoLy_MC1 * cos(skoAC1) = -0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_94, refute_0_105])).

cnf(refute_0_107, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_108, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) !=
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) |
     -skoLy_MC1 * cos(skoAC1) != -0 * cos(1047/5000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(skoBC1)),
    introduced(tautology, [equality])).

cnf(refute_0_109, plain,
    (-skoLy_MC1 * cos(skoAC1) != -0 * cos(1047/5000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(skoBC1)),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(skoBC1)),
    inference(resolve, [], [refute_0_106, refute_0_109])).

cnf(refute_0_111, plain,
    (-0 * cos(1047/5000) * cos(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) !=
     -0 * cos(1047/5000) * cos(skoBC1) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000)),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_112, plain,
    (-0 * cos(1047/5000) * cos(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_110, refute_0_111])).

cnf(refute_0_113, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_90, refute_0_112])).

cnf(refute_0_114, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_115, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1)),
    introduced(tautology, [equality])).

cnf(refute_0_116, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1)),
    inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1)),
    inference(resolve, [], [refute_0_113, refute_0_116])).

cnf(refute_0_118, plain,
    (-0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_119, plain,
    (-0 * cos(1047/5000) * cos(2443/10000) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_117, refute_0_118])).

cnf(refute_0_120, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) =
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000)),
    inference(resolve, [], [refute_0_82, refute_0_119])).

cnf(refute_0_121, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1))),
    introduced(tautology, [refl])).

cnf(refute_0_122, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) |
     -skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000))),
    introduced(tautology, [equality])).

cnf(refute_0_123, plain,
    (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
     skoLz_MC1 * cos(skoAC1) * sin(skoBC1) !=
     -0 * cos(1047/5000) * cos(2443/10000) -
     -1/4 * cos(1047/5000) * sin(2443/10000) |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000))),
    inference(resolve, [], [refute_0_121, refute_0_122])).

cnf(refute_0_124, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000))),
    inference(resolve, [], [refute_0_120, refute_0_123])).

cnf(refute_0_125, plain, (sin(skoBC1) ^ 2 = sin(skoBC1) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_126, plain,
    (sin(skoBC1) ^ 2 != sin(skoBC1) ^ 2 | sin(skoBC1) != sin(2443/10000) |
     sin(skoBC1) ^ 2 = sin(2443/10000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_127, plain,
    (sin(skoBC1) != sin(2443/10000) |
     sin(skoBC1) ^ 2 = sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_125, refute_0_126])).

cnf(refute_0_128, plain, (sin(skoBC1) ^ 2 = sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_50, refute_0_127])).

cnf(refute_0_129, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2), introduced(tautology, [refl])).

cnf(refute_0_130, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 |
     sin(skoBC1) ^ 2 != sin(2443/10000) ^ 2 |
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_131, plain,
    (sin(skoBC1) ^ 2 != sin(2443/10000) ^ 2 |
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_128, refute_0_131])).

cnf(refute_0_133, plain, (sin(skoAC1) = sin(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_134, plain,
    (sin(skoAC1) != sin(skoAC1) | skoAC1 != 1047/5000 |
     sin(skoAC1) = sin(1047/5000)), introduced(tautology, [equality])).

cnf(refute_0_135, plain,
    (skoAC1 != 1047/5000 | sin(skoAC1) = sin(1047/5000)),
    inference(resolve, [], [refute_0_133, refute_0_134])).

cnf(refute_0_136, plain, (sin(skoAC1) = sin(1047/5000)),
    inference(resolve, [], [refute_0_55, refute_0_135])).

cnf(refute_0_137, plain, (sin(skoAC1) ^ 2 = sin(skoAC1) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_138, plain,
    (sin(skoAC1) ^ 2 != sin(skoAC1) ^ 2 | sin(skoAC1) != sin(1047/5000) |
     sin(skoAC1) ^ 2 = sin(1047/5000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_139, plain,
    (sin(skoAC1) != sin(1047/5000) | sin(skoAC1) ^ 2 = sin(1047/5000) ^ 2),
    inference(resolve, [], [refute_0_137, refute_0_138])).

cnf(refute_0_140, plain, (sin(skoAC1) ^ 2 = sin(1047/5000) ^ 2),
    inference(resolve, [], [refute_0_136, refute_0_139])).

cnf(refute_0_141, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2), introduced(tautology, [refl])).

cnf(refute_0_142, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 |
     sin(skoAC1) ^ 2 != sin(1047/5000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_143, plain,
    (sin(skoAC1) ^ 2 != sin(1047/5000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2),
    inference(resolve, [], [refute_0_140, refute_0_143])).

cnf(refute_0_145, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_146, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_132, refute_0_146])).

cnf(refute_0_148, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_149, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 |
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_150, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_147, refute_0_150])).

cnf(refute_0_152, plain,
    (sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_153, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_154, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_152, refute_0_153])).

cnf(refute_0_155, plain,
    (sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_151, refute_0_154])).

cnf(refute_0_156, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_157, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_158, plain,
    (sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_156, refute_0_157])).

cnf(refute_0_159, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_155, refute_0_158])).

cnf(refute_0_160, plain, (skoTMC1 = 1357143/10000),
    inference(canonicalize, [], [normalize_0_9])).

cnf(refute_0_161, plain, (abs(skoTMC1) = abs(skoTMC1)),
    introduced(tautology, [refl])).

cnf(refute_0_162, plain,
    (abs(skoTMC1) != abs(skoTMC1) | skoTMC1 != 1357143/10000 |
     abs(skoTMC1) = abs(1357143/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_163, plain,
    (skoTMC1 != 1357143/10000 | abs(skoTMC1) = abs(1357143/10000)),
    inference(resolve, [], [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain, (abs(skoTMC1) = abs(1357143/10000)),
    inference(resolve, [], [refute_0_160, refute_0_163])).

cnf(refute_0_165, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_166, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     abs(skoTMC1) != abs(1357143/10000) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_167, plain,
    (abs(skoTMC1) != abs(1357143/10000) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    inference(resolve, [], [refute_0_164, refute_0_167])).

cnf(refute_0_169, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_170, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_168, refute_0_169])).

cnf(refute_0_171, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_159, refute_0_170])).

cnf(refute_0_172, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1))),
    introduced(tautology, [refl])).

cnf(refute_0_173, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1))),
    introduced(tautology, [equality])).

cnf(refute_0_174, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1))),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1))),
    inference(resolve, [], [refute_0_171, refute_0_174])).

cnf(refute_0_176, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000))),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_177, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000))),
    inference(resolve, [], [refute_0_175, refute_0_176])).

cnf(refute_0_178, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000))),
    inference(resolve, [], [refute_0_124, refute_0_177])).

cnf(refute_0_179, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_180, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 !=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1),
    introduced(tautology, [equality])).

cnf(refute_0_181, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1),
    inference(resolve, [], [refute_0_178, refute_0_181])).

cnf(refute_0_183, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5),
    inference(subst, [], [refute_0_38])).

cnf(refute_0_184, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - skoLz_RC1 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5),
    inference(resolve, [], [refute_0_45, refute_0_184])).

cnf(refute_0_186, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5 <
     abs(skoT_phiC1) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5 |
     abs(skoT_phiC1) <=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1),
    introduced(tautology, [equality])).

cnf(refute_0_187, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5 <
     abs(skoT_phiC1) |
     abs(skoT_phiC1) <=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (-skoLy_MC1 * cos(skoAC1) * cos(skoBC1) -
      skoLz_MC1 * cos(skoAC1) * sin(skoBC1)) - skoLz_RC1 * skoTRC1),
    inference(resolve, [], [refute_0_185, refute_0_186])).

cnf(refute_0_188, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (-0 * cos(1047/5000) * cos(2443/10000) -
      -1/4 * cos(1047/5000) * sin(2443/10000)) - 0 * 12/5 <
     abs(skoT_phiC1)),
    inference(resolve, [], [refute_0_187, refute_0_23])).

cnf(refute_0_189, plain,
    (abs(skoT_phiC1) * -1 <
     sin(2443/10000) * (cos(1047/5000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(arithmetic, [], [refute_0_188])).

cnf(refute_0_190, plain,
    (1357143/10000 < 0 | abs(1357143/10000) = 1357143/10000),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_191, plain,
    (abs(skoT_phiC1) * -1 <
     sin(2443/10000) * (cos(1047/5000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(1357143/10000) != 1357143/10000 |
     sin(2443/10000) * (cos(1047/5000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     abs(skoT_phiC1) * -1), introduced(tautology, [equality])).

cnf(refute_0_192, plain,
    (1357143/10000 < 0 |
     abs(skoT_phiC1) * -1 <
     sin(2443/10000) * (cos(1047/5000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     sin(2443/10000) * (cos(1047/5000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     abs(skoT_phiC1) * -1),
    inference(resolve, [], [refute_0_190, refute_0_191])).

cnf(refute_0_193, plain,
    (1357143/10000 < 0 |
     abs(skoT_phiC1) * -1 <
     sin(2443/10000) * (cos(1047/5000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(resolve, [], [refute_0_192, refute_0_189])).

cnf(refute_0_194, plain,
    (abs(skoT_phiC1) * -1 <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(arithmetic, [], [refute_0_193])).

cnf(refute_0_195, plain,
    (abs(skoT_phiC1) * -1 *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) |
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     abs(skoT_phiC1) * -1 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_196, plain,
    (abs(skoT_phiC1) * -1 *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(resolve, [], [refute_0_195, refute_0_194])).

cnf(refute_0_197, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     abs(skoT_phiC1) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <= -1),
    inference(arithmetic, [], [refute_0_196])).

cnf(refute_0_198, plain,
    (~ lgen(0, X_000101, sin(X_000100)) | X_000101 <= sin(X_000100)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_199, plain,
    (X_000100 < 0 | ~ lgen(0, X_000101, X_000100 - X_000100 ^ 3 / 6) |
     lgen(0, X_000101, sin(X_000100))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_200, plain,
    (X_000100 < 0 | ~ lgen(0, X_000101, X_000100 - X_000100 ^ 3 / 6) |
     X_000101 <= sin(X_000100)),
    inference(resolve, [], [refute_0_199, refute_0_198])).

cnf(refute_0_201, plain,
    (X_000100 * (1 + X_000100 * (X_000100 * -1/6)) < X_000101 |
     X_000100 < 0 | X_000101 <= sin(X_000100)),
    inference(arithmetic, [], [refute_0_200])).

cnf(refute_0_202, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     0 <= sin(2443/10000)), inference(subst, [], [refute_0_201])).

cnf(refute_0_203, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <= -1),
    introduced(tautology, [assume])).

cnf(refute_0_204, plain,
    (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) < 0 |
     0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <= -1), inference(split, [], [refute_0_203])).

cnf(refute_0_205, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(arithmetic, [], [refute_0_204])).

cnf(refute_0_206, plain,
    (12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      -12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= -1),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_207, plain,
    (sin(X_000132) < X_000133 | ~ lgen(1, sin(X_000132), X_000133)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_208, plain,
    (X_000132 < 0 |
     ~ lgen(1,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133) |
     lgen(1, sin(X_000132), X_000133)),
    inference(subst, [], [sin_upper_bound_2_pos])).

cnf(refute_0_209, plain,
    (sin(X_000132) < X_000133 | X_000132 < 0 |
     ~ lgen(1,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133)),
    inference(resolve, [], [refute_0_208, refute_0_207])).

cnf(refute_0_210, plain,
    (sin(X_000132) < X_000133 | X_000132 < 0 |
     X_000133 <=
     X_000132 *
     (1 +
      X_000132 *
      (X_000132 *
       (-1/6 +
        X_000132 *
        (X_000132 *
         (1/120 +
          X_000132 *
          (X_000132 * (-1/5040 + X_000132 * (X_000132 * 1/362880))))))))),
    inference(arithmetic, [], [refute_0_209])).

cnf(refute_0_211, plain,
    (2443/10000 < 0 |
     sin(2443/10000) <
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(subst, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     1 <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_213, plain,
    (2443/10000 < 0 |
     sin(2443/10000) <
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(subst, [], [refute_0_210])).

cnf(refute_0_214, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <= sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     -1 / (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <=
     sin(2443/10000)), inference(split, [], [refute_0_203])).

cnf(refute_0_215, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     sin(2443/10000)), inference(arithmetic, [], [refute_0_214])).

cnf(refute_0_216, plain,
    (2443/10000 < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(resolve, [], [refute_0_215, refute_0_213])).

cnf(refute_0_217, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_216])).

cnf(refute_0_218, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     1 <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_217, refute_0_212])).

cnf(refute_0_219, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= -1 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(arithmetic, [], [refute_0_218])).

cnf(refute_0_220, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      -12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= sin(2443/10000)),
    inference(split, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= sin(2443/10000)),
    inference(arithmetic, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_223, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <= 0 | sin(1047/5000) = 0),
    inference(split, [], [refute_0_222])).

cnf(refute_0_224, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_226, plain,
    (sin(1047/5000) != 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(split, [], [refute_0_225])).

cnf(refute_0_227, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_224, refute_0_226])).

cnf(refute_0_228, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= sin(2443/10000)),
    inference(resolve, [], [refute_0_221, refute_0_227])).

cnf(refute_0_229, plain,
    (2443/10000 < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(resolve, [], [refute_0_228, refute_0_211])).

cnf(refute_0_230, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      -12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= -1 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_230, refute_0_206])).

cnf(refute_0_232, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1 |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain,
    (X_000101 < sin(X_000100) | ~ lgen(1, X_000101, sin(X_000100))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_234, plain,
    (X_000100 < 0 | ~ lgen(1, X_000101, X_000100 - X_000100 ^ 3 / 6) |
     lgen(1, X_000101, sin(X_000100))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_235, plain,
    (X_000100 < 0 | X_000101 < sin(X_000100) |
     ~ lgen(1, X_000101, X_000100 - X_000100 ^ 3 / 6)),
    inference(resolve, [], [refute_0_234, refute_0_233])).

cnf(refute_0_236, plain,
    (X_000100 < 0 | X_000101 < sin(X_000100) |
     X_000100 * (1 + X_000100 * (X_000100 * -1/6)) <= X_000101),
    inference(arithmetic, [], [refute_0_235])).

cnf(refute_0_237, plain,
    (1047/5000 < 0 | 0 < sin(1047/5000) |
     1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) <= 0),
    inference(subst, [], [refute_0_236])).

cnf(refute_0_238, plain,
    (1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 <
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     sin(1047/5000) *
     -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 |
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1), inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_239, plain,
    (1047/5000 < 0 |
     sin(1047/5000) <
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880))))))))),
    inference(subst, [], [refute_0_210])).

cnf(refute_0_240, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1), introduced(tautology, [assume])).

cnf(refute_0_241, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     sin(1047/5000) *
     -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)), inference(split, [], [refute_0_240])).

cnf(refute_0_242, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000) | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_241])).

cnf(refute_0_243, plain, (0 < sin(1047/5000)),
    inference(arithmetic, [], [refute_0_237])).

cnf(refute_0_244, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)),
    inference(resolve, [], [refute_0_242, refute_0_243])).

cnf(refute_0_245, plain,
    (1047/5000 < 0 |
     -1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880))))))))),
    inference(resolve, [], [refute_0_244, refute_0_239])).

cnf(refute_0_246, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_245])).

cnf(refute_0_247, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     sin(1047/5000) *
     -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 |
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1), inference(resolve, [], [refute_0_246, refute_0_238])).

cnf(refute_0_248, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <=
     sin(1047/5000) | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (sin(X_000072) < X_000073 | ~ lgen(1, sin(X_000072), X_000073)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_250, plain,
    (X_000072 < 0 | ~ lgen(1, X_000072, X_000073) |
     lgen(1, sin(X_000072), X_000073)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_251, plain,
    (sin(X_000072) < X_000073 | X_000072 < 0 |
     ~ lgen(1, X_000072, X_000073)),
    inference(resolve, [], [refute_0_250, refute_0_249])).

cnf(refute_0_252, plain,
    (sin(X_000072) < X_000073 | X_000072 < 0 | X_000073 <= X_000072),
    inference(arithmetic, [], [refute_0_251])).

cnf(refute_0_253, plain,
    (1047/5000 < 0 |
     sin(1047/5000) <
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <=
     1047/5000), inference(subst, [], [refute_0_252])).

cnf(refute_0_254, plain,
    (sin(1047/5000) <
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467),
    inference(arithmetic, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_248, refute_0_254])).

cnf(refute_0_256, plain,
    (1047/5000 < 0 |
     -1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) <= 0),
    inference(resolve, [], [refute_0_255, refute_0_237])).

cnf(refute_0_257, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_256])).

cnf(refute_0_258, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_232, refute_0_257])).

cnf(refute_0_259, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_258])).

cnf(refute_0_260, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(1047/5000) = 0), inference(arithmetic, [], [refute_0_259])).

cnf(refute_0_261, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_260, refute_0_226])).

cnf(refute_0_262, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(resolve, [], [refute_0_261, refute_0_205])).

cnf(refute_0_263, plain,
    (0 < sin(2443/10000) * -1 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(split, [], [refute_0_262])).

cnf(refute_0_264, plain,
    (sin(2443/10000) < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(arithmetic, [], [refute_0_263])).

cnf(refute_0_265, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(resolve, [], [refute_0_202, refute_0_264])).

cnf(refute_0_266, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(arithmetic, [], [refute_0_265])).

cnf(refute_0_267, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     abs(skoT_phiC1)),
    inference(resolve, [], [refute_0_197, refute_0_266])).

cnf(refute_0_268, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(skoT_phiC1) <= 0), inference(split, [], [refute_0_267])).

cnf(refute_0_269, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1) <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_22, refute_0_268])).

cnf(refute_0_270, plain,
    (1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_269])).

cnf(refute_0_271, plain,
    (1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    introduced(tautology, [assume])).

cnf(refute_0_272, plain,
    (1 < 0 | sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) < 0 |
     0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(split, [], [refute_0_271])).

cnf(refute_0_273, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_272])).

cnf(refute_0_274, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <=
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_275, plain,
    (~ lgen(0, sin(X_000132), X_000133) | sin(X_000132) <= X_000133),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_276, plain,
    (X_000132 < 0 |
     ~ lgen(0,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133) |
     lgen(0, sin(X_000132), X_000133)),
    inference(subst, [], [sin_upper_bound_2_pos])).

cnf(refute_0_277, plain,
    (X_000132 < 0 |
     ~ lgen(0,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133) |
     sin(X_000132) <= X_000133),
    inference(resolve, [], [refute_0_276, refute_0_275])).

cnf(refute_0_278, plain,
    (X_000132 < 0 |
     X_000133 <
     X_000132 *
     (1 +
      X_000132 *
      (X_000132 *
       (-1/6 +
        X_000132 *
        (X_000132 *
         (1/120 +
          X_000132 *
          (X_000132 * (-1/5040 + X_000132 * (X_000132 * 1/362880)))))))) |
     sin(X_000132) <= X_000133),
    inference(arithmetic, [], [refute_0_277])).

cnf(refute_0_279, plain,
    (2443/10000 < 0 |
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) <=
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000))),
    inference(subst, [], [refute_0_278])).

cnf(refute_0_280, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <=
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_281, plain,
    (2443/10000 < 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) <=
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)))),
    inference(subst, [], [refute_0_278])).

cnf(refute_0_282, plain,
    (1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     sin(2443/10000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(split, [], [refute_0_271])).

cnf(refute_0_283, plain,
    (2443/10000 < 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_281, refute_0_282])).

cnf(refute_0_284, plain,
    (1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_283])).

cnf(refute_0_285, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_280, refute_0_284])).

cnf(refute_0_286, plain,
    (1 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_285])).

cnf(refute_0_287, plain,
    (1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) < sin(2443/10000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(split, [], [refute_0_286])).

cnf(refute_0_288, plain,
    (1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) < sin(2443/10000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_287, refute_0_227])).

cnf(refute_0_289, plain,
    (2443/10000 < 0 |
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_279, refute_0_288])).

cnf(refute_0_290, plain,
    (1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_289])).

cnf(refute_0_291, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_274, refute_0_290])).

cnf(refute_0_292, plain,
    (1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_291])).

cnf(refute_0_293, plain,
    (~ lgen(0, sin(X_000072), X_000073) | sin(X_000072) <= X_000073),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_294, plain,
    (X_000072 < 0 | ~ lgen(0, X_000072, X_000073) |
     lgen(0, sin(X_000072), X_000073)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_295, plain,
    (X_000072 < 0 | ~ lgen(0, X_000072, X_000073) |
     sin(X_000072) <= X_000073),
    inference(resolve, [], [refute_0_294, refute_0_293])).

cnf(refute_0_296, plain,
    (X_000073 < X_000072 | X_000072 < 0 | sin(X_000072) <= X_000073),
    inference(arithmetic, [], [refute_0_295])).

cnf(refute_0_297, plain,
    (1047/5000 < 0 |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     1047/5000 |
     sin(1047/5000) <=
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467),
    inference(subst, [], [refute_0_296])).

cnf(refute_0_298, plain,
    (1 <
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 <=
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_299, plain,
    (1047/5000 < 0 |
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880)))))))) |
     sin(1047/5000) <=
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(subst, [], [refute_0_278])).

cnf(refute_0_300, plain,
    (1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), introduced(tautology, [assume])).

cnf(refute_0_301, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     sin(1047/5000) *
     157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(split, [], [refute_0_300])).

cnf(refute_0_302, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1 | sin(1047/5000) <= 0), inference(arithmetic, [], [refute_0_301])).

cnf(refute_0_303, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_302, refute_0_243])).

cnf(refute_0_304, plain,
    (1047/5000 < 0 |
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880)))))))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_299, refute_0_303])).

cnf(refute_0_305, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(arithmetic, [], [refute_0_304])).

cnf(refute_0_306, plain,
    (1 <
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_298, refute_0_305])).

cnf(refute_0_307, plain,
    (1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1 | sin(1047/5000) <= 0), inference(arithmetic, [], [refute_0_306])).

cnf(refute_0_308, plain,
    (1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_307, refute_0_243])).

cnf(refute_0_309, plain,
    (1047/5000 < 0 |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     1047/5000 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_297, refute_0_308])).

cnf(refute_0_310, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(arithmetic, [], [refute_0_309])).

cnf(refute_0_311, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_310, refute_0_292])).

cnf(refute_0_312, plain,
    (12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_311])).

cnf(refute_0_313, plain,
    (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1 |
     sin(1047/5000) = 0), inference(arithmetic, [], [refute_0_312])).

cnf(refute_0_314, plain,
    (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_313, refute_0_226])).

cnf(refute_0_315, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_314, refute_0_273])).

cnf(refute_0_316, plain,
    (0 < sin(2443/10000) * -1 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(split, [], [refute_0_315])).

cnf(refute_0_317, plain,
    (sin(2443/10000) < 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_316])).

cnf(refute_0_318, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_202, refute_0_317])).

cnf(refute_0_319, plain,
    (sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_318])).

cnf(refute_0_320, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_319, refute_0_270])).

cnf(refute_0_321, plain, (0 <= skoT_phiC1 | abs(skoT_phiC1) = -skoT_phiC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_322, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / -skoT_phiC1 <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     abs(skoT_phiC1) != -skoT_phiC1 |
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) <=
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1)),
    introduced(tautology, [equality])).

cnf(refute_0_323, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / -skoT_phiC1 <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     0 <= skoT_phiC1 |
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) <=
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / abs(skoT_phiC1)),
    inference(resolve, [], [refute_0_321, refute_0_322])).

cnf(refute_0_324, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / -skoT_phiC1 <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     0 <= skoT_phiC1 | abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_323, refute_0_320])).

cnf(refute_0_325, plain,
    (~ lgen(0, X_000015, cos(X_000014)) | X_000015 <= cos(X_000014)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_326, plain,
    (~ lgen(0, X_000015, 1 - X_000014 ^ 2 / 2) |
     lgen(0, X_000015, cos(X_000014))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_327, plain,
    (~ lgen(0, X_000015, 1 - X_000014 ^ 2 / 2) |
     X_000015 <= cos(X_000014)),
    inference(resolve, [], [refute_0_326, refute_0_325])).

cnf(refute_0_328, plain,
    (1 + X_000014 * (X_000014 * -1/2) < X_000015 |
     X_000015 <= cos(X_000014)),
    inference(arithmetic, [], [refute_0_327])).

cnf(refute_0_329, plain,
    (1 + 1047/5000 * (1047/5000 * -1/2) <
     skoT_phiC1 *
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     skoT_phiC1 *
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     cos(1047/5000)), inference(subst, [], [refute_0_328])).

cnf(refute_0_330, plain,
    (skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_331, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_332, plain,
    (skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_333, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     sin(1047/5000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_334, plain,
    (skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_335, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_336, plain,
    (skoT_phiC1 * -1 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     0 <=
     cos(1047/5000) * -1357143/40000 +
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_337, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_9])).

cnf(refute_0_338, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     skoT_phiC1 |
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) <=
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 |
     skoT_phiC1 <= 0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_339, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1),
    inference(subst, [], [refute_0_21])).

cnf(refute_0_340, plain, (skoT_phiC1 < 0 | abs(skoT_phiC1) = skoT_phiC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_341, plain,
    (skoT_phiC1 * -1 <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(skoT_phiC1) != skoT_phiC1 |
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     abs(skoT_phiC1) * -1), introduced(tautology, [equality])).

cnf(refute_0_342, plain,
    (skoT_phiC1 < 0 |
     skoT_phiC1 * -1 <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     abs(skoT_phiC1) * -1),
    inference(resolve, [], [refute_0_340, refute_0_341])).

cnf(refute_0_343, plain,
    (skoT_phiC1 < 0 |
     skoT_phiC1 * -1 <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(resolve, [], [refute_0_342, refute_0_194])).

cnf(refute_0_344, plain,
    (skoT_phiC1 * -1 *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) |
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     skoT_phiC1 * -1 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_345, plain,
    (skoT_phiC1 < 0 |
     skoT_phiC1 * -1 *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(2443/10000) * (cos(1047/5000) * -1357143/40000) |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(resolve, [], [refute_0_344, refute_0_343])).

cnf(refute_0_346, plain,
    (skoT_phiC1 < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     skoT_phiC1 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <= -1),
    inference(arithmetic, [], [refute_0_345])).

cnf(refute_0_347, plain,
    (skoT_phiC1 < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     skoT_phiC1), inference(resolve, [], [refute_0_346, refute_0_266])).

cnf(refute_0_348, plain,
    (skoT_phiC1 < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     skoT_phiC1 <= 0), inference(split, [], [refute_0_347])).

cnf(refute_0_349, plain, (-7/200 <= skoT_phiC1),
    inference(canonicalize, [], [normalize_0_10])).

cnf(refute_0_350, plain, (skoT_phiC1 <= 367/2500),
    inference(canonicalize, [], [normalize_0_11])).

cnf(refute_0_351, plain, (0 <= skoT_phiC1 | skoT_phiC1 <= 0),
    inference(decision, [],
              [refute_0_349, refute_0_350, refute_0_55, refute_0_46,
               refute_0_95, refute_0_64, refute_0_29, refute_0_160,
               refute_0_24])).

cnf(refute_0_352, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_351, refute_0_348])).

cnf(refute_0_353, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_339, refute_0_352])).

cnf(refute_0_354, plain,
    (1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_353])).

cnf(refute_0_355, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / skoT_phiC1 <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_319, refute_0_354])).

cnf(refute_0_356, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     skoT_phiC1 | skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_338, refute_0_355])).

cnf(refute_0_357, plain,
    (skoT_phiC1 * -1 <
     sin(2443/10000) *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_356])).

cnf(refute_0_358, plain,
    (sin(2443/10000) <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     0 <=
     cos(1047/5000) * -1357143/40000 +
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     skoT_phiC1 <= 0), inference(split, [], [refute_0_357])).

cnf(refute_0_359, plain,
    (sin(2443/10000) <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <=
     cos(1047/5000) * -1357143/40000 | skoT_phiC1 <= 0),
    inference(arithmetic, [], [refute_0_358])).

cnf(refute_0_360, plain,
    (X_000029 < cos(X_000028) | ~ lgen(1, X_000029, cos(X_000028))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_361, plain,
    (~ lgen(1, X_000029,
            1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24 - X_000028 ^ 6 / 720 +
            X_000028 ^ 8 / 40320 - X_000028 ^ 10 / 3628800) |
     lgen(1, X_000029, cos(X_000028))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_362, plain,
    (X_000029 < cos(X_000028) |
     ~ lgen(1, X_000029,
            1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24 - X_000028 ^ 6 / 720 +
            X_000028 ^ 8 / 40320 - X_000028 ^ 10 / 3628800)),
    inference(resolve, [], [refute_0_361, refute_0_360])).

cnf(refute_0_363, plain,
    (X_000029 < cos(X_000028) |
     1 +
     X_000028 *
     (X_000028 *
      (-1/2 +
       X_000028 *
       (X_000028 *
        (1/24 +
         X_000028 *
         (X_000028 *
          (-1/720 +
           X_000028 *
           (X_000028 *
            (1/40320 + X_000028 * (X_000028 * -1/3628800))))))))) <=
     X_000029), inference(arithmetic, [], [refute_0_362])).

cnf(refute_0_364, plain,
    (0 < cos(1047/5000) |
     1 +
     1047/5000 *
     (1047/5000 *
      (-1/2 +
       1047/5000 *
       (1047/5000 *
        (1/24 +
         1047/5000 *
         (1047/5000 *
          (-1/720 +
           1047/5000 *
           (1047/5000 *
            (1/40320 + 1047/5000 * (1047/5000 * -1/3628800))))))))) <= 0),
    inference(subst, [], [refute_0_363])).

cnf(refute_0_365, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) < 0 |
     skoT_phiC1 < 0 | 0 < skoT_phiC1),
    inference(split, [], [refute_0_347])).

cnf(refute_0_366, plain,
    (skoT_phiC1 < 0 |
     0 < sin(2443/10000) * (cos(1047/5000) * -1357143/40000) |
     0 < skoT_phiC1), inference(arithmetic, [], [refute_0_365])).

cnf(refute_0_367, plain,
    (skoT_phiC1 < 0 |
     sin(2443/10000) < 0 / (cos(1047/5000) * -1357143/40000) |
     0 < skoT_phiC1 | 0 <= cos(1047/5000) * -1357143/40000),
    inference(split, [], [refute_0_366])).

cnf(refute_0_368, plain,
    (sin(2443/10000) < 0 | skoT_phiC1 < 0 | 0 < skoT_phiC1 |
     cos(1047/5000) <= 0), inference(arithmetic, [], [refute_0_367])).

cnf(refute_0_369, plain, (0 <= sin(2443/10000)),
    inference(arithmetic, [], [refute_0_202])).

cnf(refute_0_370, plain,
    (skoT_phiC1 < 0 | 0 < skoT_phiC1 | cos(1047/5000) <= 0),
    inference(resolve, [], [refute_0_369, refute_0_368])).

cnf(refute_0_371, plain,
    (skoT_phiC1 < 0 | 0 < skoT_phiC1 |
     1 +
     1047/5000 *
     (1047/5000 *
      (-1/2 +
       1047/5000 *
       (1047/5000 *
        (1/24 +
         1047/5000 *
         (1047/5000 *
          (-1/720 +
           1047/5000 *
           (1047/5000 *
            (1/40320 + 1047/5000 * (1047/5000 * -1/3628800))))))))) <= 0),
    inference(resolve, [], [refute_0_370, refute_0_364])).

cnf(refute_0_372, plain, (skoT_phiC1 < 0 | 0 < skoT_phiC1),
    inference(arithmetic, [], [refute_0_371])).

cnf(refute_0_373, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_374, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) * (skoT_phiC1 * -1/2) <= 0 |
     sin(1047/5000) <= 0 / (sin(1047/5000) * (skoT_phiC1 * -1/2))),
    inference(split, [], [refute_0_373])).

cnf(refute_0_375, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) * (skoT_phiC1 * -1/2) <= 0 | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_374])).

cnf(refute_0_376, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) * (skoT_phiC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_375, refute_0_243])).

cnf(refute_0_377, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1/2 <= 0 | sin(1047/5000) <= 0 / (skoT_phiC1 * -1/2)),
    inference(split, [], [refute_0_376])).

cnf(refute_0_378, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <= skoT_phiC1 | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_377])).

cnf(refute_0_379, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <= skoT_phiC1),
    inference(resolve, [], [refute_0_378, refute_0_243])).

cnf(refute_0_380, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < skoT_phiC1),
    inference(decision, [],
              [refute_0_372, refute_0_24, refute_0_160, refute_0_29,
               refute_0_64, refute_0_95, refute_0_46, refute_0_55,
               refute_0_350, refute_0_349, refute_0_379])).

cnf(refute_0_381, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(2443/10000) <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <=
     cos(1047/5000) * -1357143/40000),
    inference(resolve, [], [refute_0_359, refute_0_380])).

cnf(refute_0_382, plain,
    (X_000015 < cos(X_000014) | ~ lgen(1, X_000015, cos(X_000014))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_383, plain,
    (~ lgen(1, X_000015, 1 - X_000014 ^ 2 / 2) |
     lgen(1, X_000015, cos(X_000014))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_384, plain,
    (X_000015 < cos(X_000014) | ~ lgen(1, X_000015, 1 - X_000014 ^ 2 / 2)),
    inference(resolve, [], [refute_0_383, refute_0_382])).

cnf(refute_0_385, plain,
    (X_000015 < cos(X_000014) |
     1 + X_000014 * (X_000014 * -1/2) <= X_000015),
    inference(arithmetic, [], [refute_0_384])).

cnf(refute_0_386, plain,
    (skoT_phiC1 *
     -54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837/
     355626599934375000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <
     cos(1047/5000) |
     1 + 1047/5000 * (1047/5000 * -1/2) <=
     skoT_phiC1 *
     -54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837/
     355626599934375000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(subst, [], [refute_0_385])).

cnf(refute_0_387, plain,
    (cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(1047/5000) * -1357143/40000 |
     skoT_phiC1 *
     32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
     1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_388, plain,
    (X_000168 < sin(X_000167) | ~ lgen(1, X_000168, sin(X_000167))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_389, plain,
    (X_000167 < 0 |
     ~ lgen(1, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800) | lgen(1, X_000168, sin(X_000167))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_390, plain,
    (X_000167 < 0 | X_000168 < sin(X_000167) |
     ~ lgen(1, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800)),
    inference(resolve, [], [refute_0_389, refute_0_388])).

cnf(refute_0_391, plain,
    (X_000167 < 0 | X_000168 < sin(X_000167) |
     X_000167 *
     (1 +
      X_000167 *
      (X_000167 *
       (-1/6 +
        X_000167 *
        (X_000167 *
         (1/120 +
          X_000167 *
          (X_000167 *
           (-1/5040 +
            X_000167 *
            (X_000167 *
             (1/362880 + X_000167 * (X_000167 * -1/39916800)))))))))) <=
     X_000168), inference(arithmetic, [], [refute_0_390])).

cnf(refute_0_392, plain,
    (1047/5000 < 0 |
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(subst, [], [refute_0_391])).

cnf(refute_0_393, plain,
    (cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <=
     cos(1047/5000) * -1357143/40000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_394, plain,
    (1047/5000 < 0 |
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <
     sin(1047/5000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(subst, [], [refute_0_391])).

cnf(refute_0_395, plain,
    (cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <=
     cos(1047/5000) * -1357143/40000 |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_396, plain,
    (2443/10000 < 0 |
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <
     sin(2443/10000) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 +
              2443/10000 * (2443/10000 * -1/39916800)))))))))) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(subst, [], [refute_0_391])).

cnf(refute_0_397, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <=
     cos(1047/5000) * -1357143/40000), introduced(tautology, [assume])).

cnf(refute_0_398, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) <= 0 |
     sin(2443/10000) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(split, [], [refute_0_397])).

cnf(refute_0_399, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_400, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     sin(1047/5000) * (skoT_phiC1 * 1/2) <= 0 |
     sin(1047/5000) <= 0 / (sin(1047/5000) * (skoT_phiC1 * 1/2))),
    inference(split, [], [refute_0_399])).

cnf(refute_0_401, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     sin(1047/5000) * (skoT_phiC1 * 1/2) <= 0 | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_400])).

cnf(refute_0_402, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     sin(1047/5000) * (skoT_phiC1 * 1/2) <= 0),
    inference(resolve, [], [refute_0_401, refute_0_243])).

cnf(refute_0_403, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     skoT_phiC1 * 1/2 <= 0 | sin(1047/5000) <= 0 / (skoT_phiC1 * 1/2)),
    inference(split, [], [refute_0_402])).

cnf(refute_0_404, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     sin(1047/5000) <= 0 | skoT_phiC1 <= 0),
    inference(arithmetic, [], [refute_0_403])).

cnf(refute_0_405, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_404, refute_0_380])).

cnf(refute_0_406, plain,
    (1047/5000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) |
     1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) <= 0),
    inference(resolve, [], [refute_0_405, refute_0_237])).

cnf(refute_0_407, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))),
    inference(arithmetic, [], [refute_0_406])).

cnf(refute_0_408, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(2443/10000) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(resolve, [], [refute_0_398, refute_0_407])).

cnf(refute_0_409, plain,
    (2443/10000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 +
              2443/10000 * (2443/10000 * -1/39916800)))))))))) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(resolve, [], [refute_0_408, refute_0_396])).

cnf(refute_0_410, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(arithmetic, [], [refute_0_409])).

cnf(refute_0_411, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <=
     cos(1047/5000) * -1357143/40000 |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) <= 0),
    inference(resolve, [], [refute_0_410, refute_0_395])).

cnf(refute_0_412, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <=
     cos(1047/5000) * -1357143/40000 |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)) <= 0),
    inference(arithmetic, [], [refute_0_411])).

cnf(refute_0_413, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <=
     cos(1047/5000) * -1357143/40000),
    inference(resolve, [], [refute_0_412, refute_0_407])).

cnf(refute_0_414, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0 |
     sin(1047/5000) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(split, [], [refute_0_413])).

cnf(refute_0_415, plain,
    (sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     0 <= cos(1047/5000) * -1357143/40000),
    inference(split, [], [refute_0_413])).

cnf(refute_0_416, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     cos(1047/5000) <= 0), inference(arithmetic, [], [refute_0_415])).

cnf(refute_0_417, plain,
    (0 < cos(1047/5000) | 1 + 1047/5000 * (1047/5000 * -1/2) <= 0),
    inference(subst, [], [refute_0_385])).

cnf(refute_0_418, plain, (0 < cos(1047/5000)),
    inference(arithmetic, [], [refute_0_417])).

cnf(refute_0_419, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(resolve, [], [refute_0_416, refute_0_418])).

cnf(refute_0_420, plain,
    (1047/5000 < 0 | 1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) < 0 |
     0 <= sin(1047/5000)), inference(subst, [], [refute_0_201])).

cnf(refute_0_421, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_422, plain,
    (sin(1047/5000) <
     0 /
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     -137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(split, [], [refute_0_421])).

cnf(refute_0_423, plain,
    (sin(1047/5000) < 0 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0 |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_422])).

cnf(refute_0_424, plain,
    (0 < 0 |
     skoT_phiC1 *
     -137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 < 0 |
     0 <
     skoT_phiC1 *
     -137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(split, [], [refute_0_421])).

cnf(refute_0_425, plain,
    (skoT_phiC1 < 0 | 0 < skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_424])).

cnf(refute_0_426, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_379, refute_0_425])).

cnf(refute_0_427, plain,
    (sin(1047/5000) < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_423, refute_0_426])).

cnf(refute_0_428, plain,
    (1047/5000 < 0 | 1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_420, refute_0_427])).

cnf(refute_0_429, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_428])).

cnf(refute_0_430, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(resolve, [], [refute_0_429, refute_0_419])).

cnf(refute_0_431, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_414, refute_0_430])).

cnf(refute_0_432, plain,
    (1047/5000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_431, refute_0_394])).

cnf(refute_0_433, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     cos(1047/5000) * -1357143/40000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(arithmetic, [], [refute_0_432])).

cnf(refute_0_434, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <=
     cos(1047/5000) * -1357143/40000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_433, refute_0_393])).

cnf(refute_0_435, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(1047/5000) * -1357143/40000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      1140480000000000000000000000000000000000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_434])).

cnf(refute_0_436, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(1047/5000) * -1357143/40000),
    inference(resolve, [], [refute_0_435, refute_0_430])).

cnf(refute_0_437, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     skoT_phiC1 *
     32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
     1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) <=
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_436])).

cnf(refute_0_438, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) <=
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_437])).

cnf(refute_0_439, plain,
    (0 < 0 |
     skoT_phiC1 *
     137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     0 <
     skoT_phiC1 *
     137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000),
    inference(split, [], [refute_0_430])).

cnf(refute_0_440, plain,
    (skoT_phiC1 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     0 < skoT_phiC1), inference(arithmetic, [], [refute_0_439])).

cnf(refute_0_441, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     0 < skoT_phiC1),
    inference(resolve, [], [refute_0_379, refute_0_440])).

cnf(refute_0_442, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     sin(1047/5000) <=
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_438, refute_0_441])).

cnf(refute_0_443, plain,
    (1047/5000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_442, refute_0_392])).

cnf(refute_0_444, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     cos(1047/5000) * -1357143/40000 /
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_443])).

cnf(refute_0_445, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(1047/5000) * -1357143/40000 |
     skoT_phiC1 *
     32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
     1306800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(resolve, [], [refute_0_444, refute_0_387])).

cnf(refute_0_446, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     cos(1047/5000) <=
     skoT_phiC1 *
     -54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837/
     355626599934375000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_445])).

cnf(refute_0_447, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     cos(1047/5000) <=
     skoT_phiC1 *
     -54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837/
     355626599934375000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_446, refute_0_380])).

cnf(refute_0_448, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     1 + 1047/5000 * (1047/5000 * -1/2) <=
     skoT_phiC1 *
     -54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837/
     355626599934375000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_447, refute_0_386])).

cnf(refute_0_449, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     skoT_phiC1 <=
     -347829778344625774312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837),
    inference(arithmetic, [], [refute_0_448])).

cnf(refute_0_450, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     -347829778344625774312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837 <
     skoT_phiC1),
    inference(decision, [],
              [refute_0_372, refute_0_24, refute_0_160, refute_0_29,
               refute_0_64, refute_0_95, refute_0_46, refute_0_55,
               refute_0_350, refute_0_349, refute_0_379])).

cnf(refute_0_451, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))),
    inference(resolve, [], [refute_0_449, refute_0_450])).

cnf(refute_0_452, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(2443/10000) <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))))),
    inference(resolve, [], [refute_0_381, refute_0_451])).

cnf(refute_0_453, plain,
    (2443/10000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))))),
    inference(resolve, [], [refute_0_337, refute_0_452])).

cnf(refute_0_454, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     skoT_phiC1 * -1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))))),
    inference(arithmetic, [], [refute_0_453])).

cnf(refute_0_455, plain,
    (skoT_phiC1 * -1 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <=
     cos(1047/5000) * -1357143/40000 +
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))),
    inference(resolve, [], [refute_0_336, refute_0_454])).

cnf(refute_0_456, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) <=
     cos(1047/5000) * -1357143/40000),
    inference(arithmetic, [], [refute_0_455])).

cnf(refute_0_457, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000)))),
    inference(resolve, [], [refute_0_456, refute_0_451])).

cnf(refute_0_458, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(2443/10000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(split, [], [refute_0_457])).

cnf(refute_0_459, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(2443/10000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(arithmetic, [], [refute_0_458])).

cnf(refute_0_460, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) < 0 |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(split, [], [refute_0_457])).

cnf(refute_0_461, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     cos(1047/5000) <
     skoT_phiC1 *
     7603200000000000000000000000000000000000000000000000000/
     62396023313368661521645991742762763890535283622129834719),
    inference(arithmetic, [], [refute_0_460])).

cnf(refute_0_462, plain,
    (1 + 1047/5000 * (1047/5000 * -1/2) <
     skoT_phiC1 *
     7603200000000000000000000000000000000000000000000000000/
     62396023313368661521645991742762763890535283622129834719 |
     skoT_phiC1 *
     7603200000000000000000000000000000000000000000000000000/
     62396023313368661521645991742762763890535283622129834719 <=
     cos(1047/5000)), inference(subst, [], [refute_0_328])).

cnf(refute_0_463, plain,
    (3051402083348108529004317556175795967885084388382360411962519729/
     380160000000000000000000000000000000000000000000000000000000000 <
     skoT_phiC1 |
     skoT_phiC1 *
     7603200000000000000000000000000000000000000000000000000/
     62396023313368661521645991742762763890535283622129834719 <=
     cos(1047/5000)), inference(arithmetic, [], [refute_0_462])).

cnf(refute_0_464, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 <=
     3051402083348108529004317556175795967885084388382360411962519729/
     380160000000000000000000000000000000000000000000000000000000000),
    inference(decision, [],
              [refute_0_379, refute_0_349, refute_0_350, refute_0_55,
               refute_0_46, refute_0_95, refute_0_64, refute_0_29,
               refute_0_160, refute_0_24, refute_0_372])).

cnf(refute_0_465, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 *
     7603200000000000000000000000000000000000000000000000000/
     62396023313368661521645991742762763890535283622129834719 <=
     cos(1047/5000)),
    inference(resolve, [], [refute_0_464, refute_0_463])).

cnf(refute_0_466, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_465, refute_0_461])).

cnf(refute_0_467, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_468, plain,
    (0 <
     skoT_phiC1 *
     -137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(split, [], [refute_0_467])).

cnf(refute_0_469, plain,
    (skoT_phiC1 < 0 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(arithmetic, [], [refute_0_468])).

cnf(refute_0_470, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(resolve, [], [refute_0_379, refute_0_469])).

cnf(refute_0_471, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_470, refute_0_466])).

cnf(refute_0_472, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(2443/10000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000)))),
    inference(resolve, [], [refute_0_459, refute_0_471])).

cnf(refute_0_473, plain,
    (2443/10000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000)))),
    inference(resolve, [], [refute_0_335, refute_0_472])).

cnf(refute_0_474, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000)))),
    inference(arithmetic, [], [refute_0_473])).

cnf(refute_0_475, plain,
    (skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_334, refute_0_474])).

cnf(refute_0_476, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(arithmetic, [], [refute_0_475])).

cnf(refute_0_477, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_476, refute_0_471])).

cnf(refute_0_478, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_477])).

cnf(refute_0_479, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(arithmetic, [], [refute_0_478])).

cnf(refute_0_480, plain,
    (sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     0 |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_477])).

cnf(refute_0_481, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     cos(1047/5000) <
     skoT_phiC1 *
     7603200000000000000000000000000000000000000000000000000/
     62396023313368661521645991742762763890535283622129834719),
    inference(arithmetic, [], [refute_0_480])).

cnf(refute_0_482, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_465, refute_0_481])).

cnf(refute_0_483, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_484, plain,
    (sin(1047/5000) <
     0 /
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(split, [], [refute_0_483])).

cnf(refute_0_485, plain,
    (sin(1047/5000) < 0 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0 | skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_484])).

cnf(refute_0_486, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 20000/1357143)) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 20000/1357143)) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_487, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 20000/1357143)) |
     skoT_phiC1 * 20000/1357143 <= 0 | sin(1047/5000) = 0),
    inference(split, [], [refute_0_486])).

cnf(refute_0_488, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 20000/1357143)) |
     skoT_phiC1 <= 0 | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_487])).

cnf(refute_0_489, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 20000/1357143)) |
     sin(1047/5000) = 0),
    inference(resolve, [], [refute_0_488, refute_0_380])).

cnf(refute_0_490, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) != 0), inference(split, [], [refute_0_407])).

cnf(refute_0_491, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 20000/1357143))),
    inference(resolve, [], [refute_0_489, refute_0_490])).

cnf(refute_0_492, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) < 0 / (sin(1047/5000) * (skoT_phiC1 * 20000/1357143)) |
     0 <= sin(1047/5000) * (skoT_phiC1 * 20000/1357143)),
    inference(split, [], [refute_0_491])).

cnf(refute_0_493, plain,
    (sin(1047/5000) < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) * (skoT_phiC1 * -20000/1357143) <= 0),
    inference(arithmetic, [], [refute_0_492])).

cnf(refute_0_494, plain, (0 <= sin(1047/5000)),
    inference(arithmetic, [], [refute_0_420])).

cnf(refute_0_495, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) * (skoT_phiC1 * -20000/1357143) <= 0),
    inference(resolve, [], [refute_0_494, refute_0_493])).

cnf(refute_0_496, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <= skoT_phiC1 * -20000/1357143 |
     0 / (skoT_phiC1 * -20000/1357143) <= sin(1047/5000)),
    inference(split, [], [refute_0_495])).

cnf(refute_0_497, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <= sin(1047/5000) | skoT_phiC1 <= 0),
    inference(arithmetic, [], [refute_0_496])).

cnf(refute_0_498, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <= sin(1047/5000)),
    inference(resolve, [], [refute_0_497, refute_0_380])).

cnf(refute_0_499, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0 | skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_498, refute_0_485])).

cnf(refute_0_500, plain,
    (0 < 0 |
     skoT_phiC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <
     0 |
     0 <
     skoT_phiC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(split, [], [refute_0_483])).

cnf(refute_0_501, plain,
    (skoT_phiC1 < 0 | 0 < skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(arithmetic, [], [refute_0_500])).

cnf(refute_0_502, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 < skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(resolve, [], [refute_0_379, refute_0_501])).

cnf(refute_0_503, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(resolve, [], [refute_0_499, refute_0_502])).

cnf(refute_0_504, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_503, refute_0_482])).

cnf(refute_0_505, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_479, refute_0_504])).

cnf(refute_0_506, plain,
    (1047/5000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_333, refute_0_505])).

cnf(refute_0_507, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000))),
    inference(arithmetic, [], [refute_0_506])).

cnf(refute_0_508, plain,
    (skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <=
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_332, refute_0_507])).

cnf(refute_0_509, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0), inference(arithmetic, [], [refute_0_508])).

cnf(refute_0_510, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_509, refute_0_504])).

cnf(refute_0_511, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(split, [], [refute_0_510])).

cnf(refute_0_512, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_511])).

cnf(refute_0_513, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     sin(1047/5000) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_512, refute_0_380])).

cnf(refute_0_514, plain,
    (1047/5000 < 0 |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_331, refute_0_513])).

cnf(refute_0_515, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (skoT_phiC1 * -1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_514])).

cnf(refute_0_516, plain,
    (skoT_phiC1 * -1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     0 <=
     skoT_phiC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_330, refute_0_515])).

cnf(refute_0_517, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) <
     skoT_phiC1 *
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     skoT_phiC1 <= 0), inference(arithmetic, [], [refute_0_516])).

cnf(refute_0_518, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     cos(1047/5000) <
     skoT_phiC1 *
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_517, refute_0_380])).

cnf(refute_0_519, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     1 + 1047/5000 * (1047/5000 * -1/2) <
     skoT_phiC1 *
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_329, refute_0_518])).

cnf(refute_0_520, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     6853639791830061012926760001018098731117892354081055382793566968817062500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591 <
     skoT_phiC1), inference(arithmetic, [], [refute_0_519])).

cnf(refute_0_521, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)) |
     skoT_phiC1 <=
     6853639791830061012926760001018098731117892354081055382793566968817062500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591),
    inference(decision, [],
              [refute_0_379, refute_0_349, refute_0_350, refute_0_55,
               refute_0_46, refute_0_95, refute_0_64, refute_0_29,
               refute_0_160, refute_0_24, refute_0_372])).

cnf(refute_0_522, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))),
    inference(resolve, [], [refute_0_521, refute_0_520])).

cnf(refute_0_523, plain, (0 < skoT_phiC1 * -1/2),
    inference(split, [], [refute_0_522])).

cnf(refute_0_524, plain, (skoT_phiC1 < 0),
    inference(arithmetic, [], [refute_0_523])).

cnf(refute_0_525, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) / -skoT_phiC1 <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_324, refute_0_524])).

cnf(refute_0_526, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) /
     (skoT_phiC1 * -1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_525])).

cnf(refute_0_527, plain,
    (sin(2443/10000) * (cos(1047/5000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (skoT_phiC1 * -1) | skoT_phiC1 * -1 <= 0 | abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_17, refute_0_526])).

cnf(refute_0_528, plain,
    (skoT_phiC1 <
     sin(2443/10000) *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     0 <= skoT_phiC1 | abs(skoT_phiC1) <= 0),
    inference(arithmetic, [], [refute_0_527])).

cnf(refute_0_529, plain,
    (skoT_phiC1 <
     sin(2443/10000) *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_528, refute_0_524])).

cnf(refute_0_530, plain,
    (sin(2443/10000) <
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     0 <=
     cos(1047/5000) * -1357143/40000 +
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     abs(skoT_phiC1) <= 0), inference(split, [], [refute_0_529])).

cnf(refute_0_531, plain,
    (sin(2443/10000) <
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) <=
     cos(1047/5000) * -1357143/40000 | abs(skoT_phiC1) <= 0),
    inference(arithmetic, [], [refute_0_530])).

cnf(refute_0_532, plain,
    (-48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143))) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143))) <=
     -48903791/50000000 |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_533, plain,
    (2443/10000 < 0 |
     -48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143))) <
     sin(2443/10000) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 +
              2443/10000 * (2443/10000 * -1/39916800)))))))))) <=
     -48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)))),
    inference(subst, [], [refute_0_391])).

cnf(refute_0_534, plain,
    (X_000057 / X_000059 < cos(X_000014) |
     1 + X_000014 * (X_000014 * -1/2) <= X_000057 / X_000059),
    inference(subst, [], [refute_0_385])).

cnf(refute_0_535, plain,
    (cos(X_000014) * X_000059 < X_000057 | 0 <= X_000059 |
     cos(X_000014) <= X_000057 / X_000059),
    inference(subst, [], [leq_left_divide_mul_neg])).

cnf(refute_0_536, plain,
    (cos(X_000014) * X_000059 < X_000057 | 0 <= X_000059 |
     1 + X_000014 * (X_000014 * -1/2) <= X_000057 / X_000059),
    inference(resolve, [], [refute_0_535, refute_0_534])).

cnf(refute_0_537, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     0 <= -1357143/40000 |
     1 + 1047/5000 * (1047/5000 * -1/2) <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) /
     -1357143/40000), inference(subst, [], [refute_0_536])).

cnf(refute_0_538, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143))) <=
     -48903791/50000000), inference(arithmetic, [], [refute_0_537])).

cnf(refute_0_539, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) <=
     0 |
     sin(2443/10000) <=
     -48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)))),
    inference(split, [], [refute_0_538])).

cnf(refute_0_540, plain,
    (0 <
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_541, plain,
    (0 <
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) |
     skoT_phiC1 * -20000/1357143 <= 0 | sin(1047/5000) = 0),
    inference(split, [], [refute_0_540])).

cnf(refute_0_542, plain,
    (0 <
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) |
     0 <= skoT_phiC1 | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_541])).

cnf(refute_0_543, plain,
    (0 <
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) |
     sin(1047/5000) = 0),
    inference(resolve, [], [refute_0_542, refute_0_524])).

cnf(refute_0_544, plain, (sin(1047/5000) != 0),
    inference(split, [], [refute_0_522])).

cnf(refute_0_545, plain,
    (0 <
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143))),
    inference(resolve, [], [refute_0_543, refute_0_544])).

cnf(refute_0_546, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(2443/10000) <=
     -48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)))),
    inference(resolve, [], [refute_0_539, refute_0_545])).

cnf(refute_0_547, plain,
    (2443/10000 < 0 |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 +
              2443/10000 * (2443/10000 * -1/39916800)))))))))) <=
     -48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)))),
    inference(resolve, [], [refute_0_546, refute_0_533])).

cnf(refute_0_548, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <=
     -48903791/50000000 /
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)))),
    inference(arithmetic, [], [refute_0_547])).

cnf(refute_0_549, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143))) <=
     -48903791/50000000 |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) <=
     0), inference(resolve, [], [refute_0_548, refute_0_532])).

cnf(refute_0_550, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -20000/1357143)) <=
     0 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) <=
     -48903791/50000000), inference(arithmetic, [], [refute_0_549])).

cnf(refute_0_551, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) <=
     -48903791/50000000),
    inference(resolve, [], [refute_0_550, refute_0_545])).

cnf(refute_0_552, plain,
    (sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) < 0 |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     0 <= -48903791/50000000), inference(split, [], [refute_0_551])).

cnf(refute_0_553, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) |
     0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))),
    inference(arithmetic, [], [refute_0_552])).

cnf(refute_0_554, plain,
    (-48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -48903791/50000000 |
     skoT_phiC1 *
     -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
     44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_555, plain,
    (1047/5000 < 0 |
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(subst, [], [refute_0_391])).

cnf(refute_0_556, plain,
    (-48903791/50000000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) <=
     -48903791/50000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_557, plain,
    (1047/5000 < 0 |
     -48903791/50000000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) <
     sin(1047/5000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     -48903791/50000000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000))),
    inference(subst, [], [refute_0_391])).

cnf(refute_0_558, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0 |
     sin(1047/5000) <=
     -48903791/50000000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000))),
    inference(split, [], [refute_0_551])).

cnf(refute_0_559, plain,
    (1047/5000 < 0 |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     -48903791/50000000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_558, refute_0_557])).

cnf(refute_0_560, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     -48903791/50000000 /
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_559])).

cnf(refute_0_561, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       38694861216000000000000000000000000000000000000000000)) <=
     -48903791/50000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_560, refute_0_556])).

cnf(refute_0_562, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -48903791/50000000), inference(arithmetic, [], [refute_0_561])).

cnf(refute_0_563, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0 |
     skoT_phiC1 *
     -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
     44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) <=
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_562])).

cnf(refute_0_564, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     0 <= skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0 |
     sin(1047/5000) <=
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_563])).

cnf(refute_0_565, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0 |
     sin(1047/5000) <=
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_564, refute_0_524])).

cnf(refute_0_566, plain,
    (1047/5000 < 0 |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_565, refute_0_555])).

cnf(refute_0_567, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     -48903791/50000000 /
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_566])).

cnf(refute_0_568, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
      44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -48903791/50000000 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0 |
     skoT_phiC1 *
     -32852784343073059396093857870719729730718633900015245047878296561828718349789235022087700498903/
     44337861810000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(resolve, [], [refute_0_567, refute_0_554])).

cnf(refute_0_569, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     0 <= skoT_phiC1 |
     347829778344625774312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837 <=
     skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(arithmetic, [], [refute_0_568])).

cnf(refute_0_570, plain,
    (skoT_phiC1 < 0 |
     347829778344625774312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837 <=
     skoT_phiC1),
    inference(decision, [],
              [refute_0_524, refute_0_372, refute_0_24, refute_0_160,
               refute_0_29, refute_0_64, refute_0_95, refute_0_46,
               refute_0_55, refute_0_350, refute_0_349])).

cnf(refute_0_571, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     347829778344625774312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837 <=
     skoT_phiC1 |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_569, refute_0_570])).

cnf(refute_0_572, plain,
    (skoT_phiC1 <
     347829778344625774312500000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     54775941404623812325157358806122888772015033052880987685171194475094192125070137839834750576270649801209958069111508772915428594859318132837),
    inference(decision, [],
              [refute_0_524, refute_0_372, refute_0_24, refute_0_160,
               refute_0_29, refute_0_64, refute_0_95, refute_0_46,
               refute_0_55, refute_0_350, refute_0_349])).

cnf(refute_0_573, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) <= 0),
    inference(resolve, [], [refute_0_571, refute_0_572])).

cnf(refute_0_574, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))),
    inference(resolve, [], [refute_0_573, refute_0_553])).

cnf(refute_0_575, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     sin(1047/5000) <
     0 /
     (skoT_phiC1 *
      137928037016074197461091406895432752238788286029099/
      38694861216000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     137928037016074197461091406895432752238788286029099/
     38694861216000000000000000000000000000000000000000000),
    inference(split, [], [refute_0_574])).

cnf(refute_0_576, plain,
    (sin(1047/5000) < 0 |
     cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     0 <= skoT_phiC1), inference(arithmetic, [], [refute_0_575])).

cnf(refute_0_577, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) |
     0 <= skoT_phiC1),
    inference(resolve, [], [refute_0_494, refute_0_576])).

cnf(refute_0_578, plain,
    (cos(1047/5000) * -1357143/40000 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2)))),
    inference(resolve, [], [refute_0_577, refute_0_524])).

cnf(refute_0_579, plain,
    (sin(2443/10000) <
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_531, refute_0_578])).

cnf(refute_0_580, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_16, refute_0_579])).

cnf(refute_0_581, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     skoT_phiC1 /
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_580])).

cnf(refute_0_582, plain,
    (skoT_phiC1 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (cos(1047/5000) * -1357143/40000 +
      sin(2443/10000) *
      (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2)))) |
     0 <=
     cos(1047/5000) * -1357143/40000 +
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * 1/2))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_15, refute_0_581])).

cnf(refute_0_583, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (skoT_phiC1 * -1/2))) <=
     cos(1047/5000) * -1357143/40000 | abs(skoT_phiC1) <= 0),
    inference(arithmetic, [], [refute_0_582])).

cnf(refute_0_584, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_583, refute_0_578])).

cnf(refute_0_585, plain,
    (sin(2443/10000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     abs(skoT_phiC1) <= 0), inference(split, [], [refute_0_584])).

cnf(refute_0_586, plain,
    (sin(2443/10000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_585])).

cnf(refute_0_587, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     skoT_phiC1 *
     -137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_467])).

cnf(refute_0_588, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <= skoT_phiC1 | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_587])).

cnf(refute_0_589, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) = 0),
    inference(resolve, [], [refute_0_588, refute_0_524])).

cnf(refute_0_590, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000))),
    inference(resolve, [], [refute_0_589, refute_0_544])).

cnf(refute_0_591, plain,
    (sin(2443/10000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_586, refute_0_590])).

cnf(refute_0_592, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_14, refute_0_591])).

cnf(refute_0_593, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_592])).

cnf(refute_0_594, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (skoT_phiC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_13, refute_0_593])).

cnf(refute_0_595, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_594])).

cnf(refute_0_596, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_595, refute_0_590])).

cnf(refute_0_597, plain,
    (sin(1047/5000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     abs(skoT_phiC1) <= 0), inference(split, [], [refute_0_596])).

cnf(refute_0_598, plain,
    (sin(1047/5000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0 | abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_597])).

cnf(refute_0_599, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     skoT_phiC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) <=
     0 /
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_483])).

cnf(refute_0_600, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_phiC1 | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_599])).

cnf(refute_0_601, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_600, refute_0_524])).

cnf(refute_0_602, plain,
    (0 <
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_601, refute_0_243])).

cnf(refute_0_603, plain,
    (sin(1047/5000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_598, refute_0_602])).

cnf(refute_0_604, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_12, refute_0_603])).

cnf(refute_0_605, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_604])).

cnf(refute_0_606, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (skoT_phiC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (skoT_phiC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_11, refute_0_605])).

cnf(refute_0_607, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (skoT_phiC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     0 | abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_606])).

cnf(refute_0_608, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_607, refute_0_602])).

cnf(refute_0_609, plain,
    (sin(1047/5000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     abs(skoT_phiC1) <= 0), inference(split, [], [refute_0_608])).

cnf(refute_0_610, plain,
    (sin(1047/5000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_phiC1 | abs(skoT_phiC1) <= 0),
    inference(arithmetic, [], [refute_0_609])).

cnf(refute_0_611, plain,
    (sin(1047/5000) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_610, refute_0_524])).

cnf(refute_0_612, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_10, refute_0_611])).

cnf(refute_0_613, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (skoT_phiC1 +
      cos(1047/5000) *
      62396023313368661521645991742762763890535283622129834719/
      7603200000000000000000000000000000000000000000000000000) /
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_612])).

cnf(refute_0_614, plain,
    (skoT_phiC1 +
     cos(1047/5000) *
     62396023313368661521645991742762763890535283622129834719/
     7603200000000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (skoT_phiC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     skoT_phiC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_5, refute_0_613])).

cnf(refute_0_615, plain,
    (cos(1047/5000) <
     skoT_phiC1 *
     -852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     0 <= skoT_phiC1 | abs(skoT_phiC1) <= 0),
    inference(arithmetic, [], [refute_0_614])).

cnf(refute_0_616, plain,
    (cos(1047/5000) <
     skoT_phiC1 *
     -852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_615, refute_0_524])).

cnf(refute_0_617, plain,
    (1 +
     1047/5000 *
     (1047/5000 *
      (-1/2 +
       1047/5000 *
       (1047/5000 *
        (1/24 +
         1047/5000 *
         (1047/5000 *
          (-1/720 +
           1047/5000 *
           (1047/5000 *
            (1/40320 + 1047/5000 * (1047/5000 * -1/3628800))))))))) <
     skoT_phiC1 *
     -852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591/
     7007268405664154965948100016436454518544270273526499594677196875000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_4, refute_0_616])).

cnf(refute_0_618, plain,
    (skoT_phiC1 <
     -6854200335540288400887629020185495225546766341663665207726298073178119574074858634786285021699173717950000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591 |
     abs(skoT_phiC1) <= 0), inference(arithmetic, [], [refute_0_617])).

cnf(refute_0_619, plain,
    (-6854200335540288400887629020185495225546766341663665207726298073178119574074858634786285021699173717950000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     852783814560907533649182792003897345490708212447763745126482630890320494776842328795816959609528183662730447746358143961165203669993489502219426434965371134353429485963336068260760164895796591 <=
     skoT_phiC1),
    inference(decision, [],
              [refute_0_349, refute_0_350, refute_0_55, refute_0_46,
               refute_0_95, refute_0_64, refute_0_29, refute_0_160,
               refute_0_24, refute_0_372, refute_0_524])).

cnf(refute_0_620, plain, (abs(skoT_phiC1) <= 0),
    inference(resolve, [], [refute_0_619, refute_0_618])).

cnf(refute_0_621, plain,
    (0 < abs(skoT_phiC1) | abs(skoT_phiC1) != -skoT_phiC1 |
     -skoT_phiC1 <= 0), introduced(tautology, [equality])).

cnf(refute_0_622, plain,
    (0 < abs(skoT_phiC1) | 0 <= skoT_phiC1 | -skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_321, refute_0_621])).

cnf(refute_0_623, plain, (0 <= skoT_phiC1 | -skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_620, refute_0_622])).

cnf(refute_0_624, plain, (-skoT_phiC1 <= 0),
    inference(resolve, [], [refute_0_623, refute_0_524])).

cnf(refute_0_625, plain, (0 <= skoT_phiC1),
    inference(arithmetic, [], [refute_0_624])).

cnf(refute_0_626, plain, ($false),
    inference(resolve, [], [refute_0_625, refute_0_524])).
SZS output end CNFRefutation for /home/omar/Desktop/Omar - ex/My Problems/Helicopter/Torque_phi.tptp

Processor time: 51.312 = 38.456 (Metis) + 12.856 (RCF)
Maximum weight in proof search: 618


