%--------------------------------------------------------------------------
% Proof Results of "Torque_theta.tptp"
%--------------------------------------------------------------------------
Z3 version 4.4.2 - 64 bit
---------------------------------------------------------------------------
SZS status Theorem for /home/omar/Desktop/Omar - ex/My Problems/Helicopter/Torque_theta.tptp

SZS output start CNFRefutation for /home/omar/Desktop/Omar - ex/My Problems/Helicopter/Torque_theta.tptp
cnf(lgen_le_neg, axiom, (X <= Y | ~ lgen(0, X, Y))).

cnf(lgen_less_neg, axiom, (X < Y | ~ lgen(1, X, Y))).

cnf(leq_left_divide_mul_pos, axiom, (Y * Z < X | X / Z <= Y | Z <= 0)).

cnf(leq_left_mul_divide_pos, axiom, (Y / Z < X | X * Z <= Y | Z <= 0)).

cnf(leq_right_divide_mul_pos, axiom, (Y < X * Z | X <= Y / Z | Z <= 0)).

cnf(leq_right_mul_divide_pos, axiom, (Y < X / Z | X <= Y * Z | Z <= 0)).

cnf(leq_right_divide_mul_neg, axiom, (Y < X * Z | Y / Z <= X | 0 <= Z)).

cnf(leq_right_mul_divide_neg, axiom, (Y < X / Z | Y * Z <= X | 0 <= Z)).

cnf(sqrt_upper_bound_0, axiom,
    (X < 0 | ~ lgen(R, (X + 1) / 2, Y) | lgen(R, sqrt(X), Y))).

cnf(abs_nonnegative, axiom, (X < 0 | abs(X) = X)).

cnf(abs_negative, axiom, (0 <= X | abs(X) = -X)).

cnf(sin_upper_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, X, Y) | lgen(R, sin(X), Y))).

cnf(sin_upper_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880,
            Y) | lgen(R, sin(X), Y))).

cnf(sin_lower_bound_0_pos, axiom,
    (X < 0 | ~ lgen(R, Y, X - X ^ 3 / 6) | lgen(R, Y, sin(X)))).

cnf(sin_lower_bound_2_pos, axiom,
    (X < 0 |
     ~ lgen(R, Y,
            X - X ^ 3 / 6 + X ^ 5 / 120 - X ^ 7 / 5040 + X ^ 9 / 362880 -
            X ^ 11 / 39916800) | lgen(R, Y, sin(X)))).

cnf(cos_lower_bound_0, axiom,
    (~ lgen(R, Y, 1 - X ^ 2 / 2) | lgen(R, Y, cos(X)))).

cnf(cos_lower_bound_2, axiom,
    (~ lgen(R, Y,
            1 - X ^ 2 / 2 + X ^ 4 / 24 - X ^ 6 / 720 + X ^ 8 / 40320 -
            X ^ 10 / 3628800) | lgen(R, Y, cos(X)))).

fof('Torque_theta', conjecture,
    (! [T_theta, A, B, TM, TR, Lx_M, Lz_M] :
       ((-1061/10000 <= T_theta & T_theta <= 12667/10000 &
         TM = 1357143/10000 & TR = 12/5 & A = 1047/5000 & B = 2443/10000 &
         Lx_M = 0 & Lz_M = -1/4) =>
        abs(T_theta) <=
        abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
        (Lx_M * cos(A) * cos(B) + 1/4 * sin(A) * cos(B)) - 1/50 * TR))).

fof(subgoal_0, plain,
    (! [T_theta, A, B, TM, TR, Lx_M, Lz_M] :
       ((-1061/10000 <= T_theta & T_theta <= 12667/10000 &
         TM = 1357143/10000 & TR = 12/5 & A = 1047/5000 & B = 2443/10000 &
         Lx_M = 0 & Lz_M = -1/4) =>
        abs(T_theta) <=
        abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
        (Lx_M * cos(A) * cos(B) + 1/4 * sin(A) * cos(B)) - 1/50 * TR)),
    inference(strip, [], ['Torque_theta'])).

fof(negate_0_0, plain,
    (~ ! [T_theta, A, B, TM, TR, Lx_M, Lz_M] :
         ((-1061/10000 <= T_theta & T_theta <= 12667/10000 &
           TM = 1357143/10000 & TR = 12/5 & A = 1047/5000 &
           B = 2443/10000 & Lx_M = 0 & Lz_M = -1/4) =>
          abs(T_theta) <=
          abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
          (Lx_M * cos(A) * cos(B) + 1/4 * sin(A) * cos(B)) - 1/50 * TR)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (? [A, B, Lx_M, TM, TR, T_theta] :
       (abs(TM) / sqrt(1 - sin(A) ^ 2 * sin(B) ^ 2) *
        (Lx_M * cos(A) * cos(B) + 1/4 * sin(A) * cos(B)) - 1/50 * TR <
        abs(T_theta) & T_theta <= 12667/10000 & -1061/10000 <= T_theta &
        A = 1047/5000 & B = 2443/10000 & Lx_M = 0 & TM = 1357143/10000 &
        TR = 12/5 & ? [Lz_M] : Lz_M = -1/4)),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 <
     abs(skoT_thetaC1) & -1061/10000 <= skoT_thetaC1 &
     skoT_thetaC1 <= 12667/10000 & skoAC1 = 1047/5000 &
     skoBC1 = 2443/10000 & skoLx_MC1 = 0 & skoTMC1 = 1357143/10000 &
     skoTRC1 = 12/5 & ? [Lz_M] : Lz_M = -1/4),
    inference(skolemize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 <
     abs(skoT_thetaC1)), inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_3, plain, (skoTRC1 = 12/5),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_4, plain, (skoBC1 = 2443/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_5, plain, (skoAC1 = 1047/5000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_6, plain, (skoLx_MC1 = 0),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_7, plain, (skoTMC1 = 1357143/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_8, plain, (-1061/10000 <= skoT_thetaC1),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_9, plain, (skoT_thetaC1 <= 12667/10000),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_10, plain, (? [Lz_M] : Lz_M = -1/4),
    inference(conjunct, [], [normalize_0_1])).

fof(normalize_0_11, plain, (skoLz_MC1 = -1/4),
    inference(skolemize, [], [normalize_0_10])).

cnf(refute_0_0, plain,
    (X_000101 < sin(X_000100) | ~ lgen(1, X_000101, sin(X_000100))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_1, plain,
    (X_000100 < 0 | ~ lgen(1, X_000101, X_000100 - X_000100 ^ 3 / 6) |
     lgen(1, X_000101, sin(X_000100))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_2, plain,
    (X_000100 < 0 | X_000101 < sin(X_000100) |
     ~ lgen(1, X_000101, X_000100 - X_000100 ^ 3 / 6)),
    inference(resolve, [], [refute_0_1, refute_0_0])).

cnf(refute_0_3, plain,
    (X_000100 < 0 | X_000101 < sin(X_000100) |
     X_000100 * (1 + X_000100 * (X_000100 * -1/6)) <= X_000101),
    inference(arithmetic, [], [refute_0_2])).

cnf(refute_0_4, plain,
    (1047/5000 < 0 | 0 < sin(1047/5000) |
     1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) <= 0),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_5, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_6, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) <= 0 / (sin(1047/5000) * sin(1047/5000))),
    inference(split, [], [refute_0_5])).

cnf(refute_0_7, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 | sin(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_6])).

cnf(refute_0_8, plain,
    (2443/10000 < 0 | 0 < sin(2443/10000) |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) <= 0),
    inference(subst, [], [refute_0_3])).

cnf(refute_0_9, plain, (0 < sin(2443/10000)),
    inference(arithmetic, [], [refute_0_8])).

cnf(refute_0_10, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_7, refute_0_9])).

cnf(refute_0_11, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(1047/5000) <= 0 | sin(1047/5000) <= 0 / sin(1047/5000)),
    inference(split, [], [refute_0_10])).

cnf(refute_0_12, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     sin(1047/5000) <= 0), inference(arithmetic, [], [refute_0_11])).

cnf(refute_0_13, plain,
    (1047/5000 < 0 |
     0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) <= 0),
    inference(resolve, [], [refute_0_12, refute_0_4])).

cnf(refute_0_14, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))),
    inference(arithmetic, [], [refute_0_13])).

cnf(refute_0_15, plain,
    (0 < 0 | sin(1047/5000) * sin(1047/5000) < 0 |
     0 < sin(1047/5000) * sin(1047/5000)),
    inference(split, [], [refute_0_14])).

cnf(refute_0_16, plain,
    (0 < sin(1047/5000) * (sin(1047/5000) * -1) |
     0 < sin(1047/5000) * sin(1047/5000)),
    inference(arithmetic, [], [refute_0_15])).

cnf(refute_0_17, plain,
    (~ lgen(0, X_000015, cos(X_000014)) | X_000015 <= cos(X_000014)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_18, plain,
    (~ lgen(0, X_000015, 1 - X_000014 ^ 2 / 2) |
     lgen(0, X_000015, cos(X_000014))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_19, plain,
    (~ lgen(0, X_000015, 1 - X_000014 ^ 2 / 2) |
     X_000015 <= cos(X_000014)),
    inference(resolve, [], [refute_0_18, refute_0_17])).

cnf(refute_0_20, plain,
    (1 + X_000014 * (X_000014 * -1/2) < X_000015 |
     X_000015 <= cos(X_000014)), inference(arithmetic, [], [refute_0_19])).

cnf(refute_0_21, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     cos(2443/10000)), inference(subst, [], [refute_0_20])).

cnf(refute_0_22, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_23, plain,
    (~ lgen(0, X_000168, sin(X_000167)) | X_000168 <= sin(X_000167)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_24, plain,
    (X_000167 < 0 |
     ~ lgen(0, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800) | lgen(0, X_000168, sin(X_000167))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_25, plain,
    (X_000167 < 0 |
     ~ lgen(0, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800) | X_000168 <= sin(X_000167)),
    inference(resolve, [], [refute_0_24, refute_0_23])).

cnf(refute_0_26, plain,
    (X_000167 *
     (1 +
      X_000167 *
      (X_000167 *
       (-1/6 +
        X_000167 *
        (X_000167 *
         (1/120 +
          X_000167 *
          (X_000167 *
           (-1/5040 +
            X_000167 *
            (X_000167 *
             (1/362880 + X_000167 * (X_000167 * -1/39916800)))))))))) <
     X_000168 | X_000167 < 0 | X_000168 <= sin(X_000167)),
    inference(arithmetic, [], [refute_0_25])).

cnf(refute_0_27, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_28, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_29, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     sin(1047/5000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_30, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_31, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_32, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) |
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_33, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_34, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (6/125 + skoT_thetaC1) | 6/125 + skoT_thetaC1 <= 0 |
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) <=
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_35, plain,
    (~ lgen(0, sqrt(X_000086), X_000087) | sqrt(X_000086) <= X_000087),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_36, plain,
    (X_000086 < 0 | ~ lgen(0, (X_000086 + 1) / 2, X_000087) |
     lgen(0, sqrt(X_000086), X_000087)),
    inference(subst, [], [sqrt_upper_bound_0])).

cnf(refute_0_37, plain,
    (X_000086 < 0 | ~ lgen(0, (X_000086 + 1) / 2, X_000087) |
     sqrt(X_000086) <= X_000087),
    inference(resolve, [], [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (X_000086 < 0 | X_000087 < 1/2 + X_000086 * 1/2 |
     sqrt(X_000086) <= X_000087),
    inference(arithmetic, [], [refute_0_37])).

cnf(refute_0_39, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1)), inference(subst, [], [refute_0_38])).

cnf(refute_0_40, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 <
     abs(skoT_thetaC1)), inference(canonicalize, [], [normalize_0_2])).

cnf(refute_0_41, plain, (skoTRC1 = 12/5),
    inference(canonicalize, [], [normalize_0_3])).

cnf(refute_0_42, plain, (1/50 * skoTRC1 = 1/50 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_43, plain,
    (1/50 * skoTRC1 != 1/50 * skoTRC1 | skoTRC1 != 12/5 |
     1/50 * skoTRC1 = 1/50 * 12/5), introduced(tautology, [equality])).

cnf(refute_0_44, plain, (skoTRC1 != 12/5 | 1/50 * skoTRC1 = 1/50 * 12/5),
    inference(resolve, [], [refute_0_42, refute_0_43])).

cnf(refute_0_45, plain, (1/50 * skoTRC1 = 1/50 * 12/5),
    inference(resolve, [], [refute_0_41, refute_0_44])).

cnf(refute_0_46, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_47, plain,
    (1/50 * skoTRC1 != 1/50 * 12/5 |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5),
    introduced(tautology, [equality])).

cnf(refute_0_48, plain,
    (1/50 * skoTRC1 != 1/50 * 12/5 |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5),
    inference(resolve, [], [refute_0_46, refute_0_47])).

cnf(refute_0_49, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5),
    inference(resolve, [], [refute_0_45, refute_0_48])).

cnf(refute_0_50, plain, (skoBC1 = 2443/10000),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_51, plain, (cos(skoBC1) = cos(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_52, plain,
    (cos(skoBC1) != cos(skoBC1) | skoBC1 != 2443/10000 |
     cos(skoBC1) = cos(2443/10000)), introduced(tautology, [equality])).

cnf(refute_0_53, plain,
    (skoBC1 != 2443/10000 | cos(skoBC1) = cos(2443/10000)),
    inference(resolve, [], [refute_0_51, refute_0_52])).

cnf(refute_0_54, plain, (cos(skoBC1) = cos(2443/10000)),
    inference(resolve, [], [refute_0_50, refute_0_53])).

cnf(refute_0_55, plain,
    (1/4 * sin(1047/5000) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(skoBC1)), introduced(tautology, [refl])).

cnf(refute_0_56, plain,
    (1/4 * sin(1047/5000) * cos(skoBC1) !=
     1/4 * sin(1047/5000) * cos(skoBC1) | cos(skoBC1) != cos(2443/10000) |
     1/4 * sin(1047/5000) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(2443/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_57, plain,
    (cos(skoBC1) != cos(2443/10000) |
     1/4 * sin(1047/5000) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_55, refute_0_56])).

cnf(refute_0_58, plain,
    (1/4 * sin(1047/5000) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_54, refute_0_57])).

cnf(refute_0_59, plain, (skoAC1 = 1047/5000),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_60, plain, (sin(skoAC1) = sin(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_61, plain,
    (sin(skoAC1) != sin(skoAC1) | skoAC1 != 1047/5000 |
     sin(skoAC1) = sin(1047/5000)), introduced(tautology, [equality])).

cnf(refute_0_62, plain,
    (skoAC1 != 1047/5000 | sin(skoAC1) = sin(1047/5000)),
    inference(resolve, [], [refute_0_60, refute_0_61])).

cnf(refute_0_63, plain, (sin(skoAC1) = sin(1047/5000)),
    inference(resolve, [], [refute_0_59, refute_0_62])).

cnf(refute_0_64, plain, (1/4 * sin(skoAC1) = 1/4 * sin(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_65, plain,
    (1/4 * sin(skoAC1) != 1/4 * sin(skoAC1) |
     sin(skoAC1) != sin(1047/5000) |
     1/4 * sin(skoAC1) = 1/4 * sin(1047/5000)),
    introduced(tautology, [equality])).

cnf(refute_0_66, plain,
    (sin(skoAC1) != sin(1047/5000) |
     1/4 * sin(skoAC1) = 1/4 * sin(1047/5000)),
    inference(resolve, [], [refute_0_64, refute_0_65])).

cnf(refute_0_67, plain, (1/4 * sin(skoAC1) = 1/4 * sin(1047/5000)),
    inference(resolve, [], [refute_0_63, refute_0_66])).

cnf(refute_0_68, plain,
    (1/4 * sin(skoAC1) * cos(skoBC1) = 1/4 * sin(skoAC1) * cos(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_69, plain,
    (1/4 * sin(skoAC1) != 1/4 * sin(1047/5000) |
     1/4 * sin(skoAC1) * cos(skoBC1) != 1/4 * sin(skoAC1) * cos(skoBC1) |
     1/4 * sin(skoAC1) * cos(skoBC1) = 1/4 * sin(1047/5000) * cos(skoBC1)),
    introduced(tautology, [equality])).

cnf(refute_0_70, plain,
    (1/4 * sin(skoAC1) != 1/4 * sin(1047/5000) |
     1/4 * sin(skoAC1) * cos(skoBC1) = 1/4 * sin(1047/5000) * cos(skoBC1)),
    inference(resolve, [], [refute_0_68, refute_0_69])).

cnf(refute_0_71, plain,
    (1/4 * sin(skoAC1) * cos(skoBC1) = 1/4 * sin(1047/5000) * cos(skoBC1)),
    inference(resolve, [], [refute_0_67, refute_0_70])).

cnf(refute_0_72, plain, (X0 = X0), introduced(tautology, [refl])).

cnf(refute_0_73, plain, (X0 != X0 | X0 != Y0 | Y0 = X0),
    introduced(tautology, [equality])).

cnf(refute_0_74, plain, (X0 != Y0 | Y0 = X0),
    inference(resolve, [], [refute_0_72, refute_0_73])).

cnf(refute_0_75, plain, (Y0 != X0 | Y0 != Z0 | X0 = Z0),
    introduced(tautology, [equality])).

cnf(refute_0_76, plain, (X0 != Y0 | Y0 != Z0 | X0 = Z0),
    inference(resolve, [], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (1/4 * sin(1047/5000) * cos(skoBC1) !=
     1/4 * sin(1047/5000) * cos(2443/10000) |
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     1/4 * sin(1047/5000) * cos(skoBC1) |
     1/4 * sin(skoAC1) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_78, plain,
    (1/4 * sin(1047/5000) * cos(skoBC1) !=
     1/4 * sin(1047/5000) * cos(2443/10000) |
     1/4 * sin(skoAC1) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_71, refute_0_77])).

cnf(refute_0_79, plain,
    (1/4 * sin(skoAC1) * cos(skoBC1) =
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_58, refute_0_78])).

cnf(refute_0_80, plain,
    (0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1)), introduced(tautology, [refl])).

cnf(refute_0_81, plain,
    (1/4 * sin(skoAC1) * cos(skoBC1) !=
     1/4 * sin(1047/5000) * cos(2443/10000) |
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) |
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_82, plain,
    (1/4 * sin(skoAC1) * cos(skoBC1) !=
     1/4 * sin(1047/5000) * cos(2443/10000) |
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_80, refute_0_81])).

cnf(refute_0_83, plain,
    (0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_79, refute_0_82])).

cnf(refute_0_84, plain,
    (0 * cos(1047/5000) * cos(skoBC1) = 0 * cos(1047/5000) * cos(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_85, plain,
    (0 * cos(1047/5000) * cos(skoBC1) != 0 * cos(1047/5000) * cos(skoBC1) |
     cos(skoBC1) != cos(2443/10000) |
     0 * cos(1047/5000) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_86, plain,
    (cos(skoBC1) != cos(2443/10000) |
     0 * cos(1047/5000) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_84, refute_0_85])).

cnf(refute_0_87, plain,
    (0 * cos(1047/5000) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_54, refute_0_86])).

cnf(refute_0_88, plain, (cos(skoAC1) = cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_89, plain,
    (cos(skoAC1) != cos(skoAC1) | skoAC1 != 1047/5000 |
     cos(skoAC1) = cos(1047/5000)), introduced(tautology, [equality])).

cnf(refute_0_90, plain,
    (skoAC1 != 1047/5000 | cos(skoAC1) = cos(1047/5000)),
    inference(resolve, [], [refute_0_88, refute_0_89])).

cnf(refute_0_91, plain, (cos(skoAC1) = cos(1047/5000)),
    inference(resolve, [], [refute_0_59, refute_0_90])).

cnf(refute_0_92, plain, (0 * cos(skoAC1) = 0 * cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_93, plain,
    (0 * cos(skoAC1) != 0 * cos(skoAC1) | cos(skoAC1) != cos(1047/5000) |
     0 * cos(skoAC1) = 0 * cos(1047/5000)),
    introduced(tautology, [equality])).

cnf(refute_0_94, plain,
    (cos(skoAC1) != cos(1047/5000) | 0 * cos(skoAC1) = 0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_92, refute_0_93])).

cnf(refute_0_95, plain, (0 * cos(skoAC1) = 0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_91, refute_0_94])).

cnf(refute_0_96, plain, (skoLx_MC1 = 0),
    inference(canonicalize, [], [normalize_0_6])).

cnf(refute_0_97, plain,
    (skoLx_MC1 * cos(skoAC1) = skoLx_MC1 * cos(skoAC1)),
    introduced(tautology, [refl])).

cnf(refute_0_98, plain,
    (skoLx_MC1 * cos(skoAC1) != skoLx_MC1 * cos(skoAC1) | skoLx_MC1 != 0 |
     skoLx_MC1 * cos(skoAC1) = 0 * cos(skoAC1)),
    introduced(tautology, [equality])).

cnf(refute_0_99, plain,
    (skoLx_MC1 != 0 | skoLx_MC1 * cos(skoAC1) = 0 * cos(skoAC1)),
    inference(resolve, [], [refute_0_97, refute_0_98])).

cnf(refute_0_100, plain, (skoLx_MC1 * cos(skoAC1) = 0 * cos(skoAC1)),
    inference(resolve, [], [refute_0_96, refute_0_99])).

cnf(refute_0_101, plain,
    (0 * cos(skoAC1) != 0 * cos(1047/5000) |
     skoLx_MC1 * cos(skoAC1) != 0 * cos(skoAC1) |
     skoLx_MC1 * cos(skoAC1) = 0 * cos(1047/5000)),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_102, plain,
    (0 * cos(skoAC1) != 0 * cos(1047/5000) |
     skoLx_MC1 * cos(skoAC1) = 0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_100, refute_0_101])).

cnf(refute_0_103, plain, (skoLx_MC1 * cos(skoAC1) = 0 * cos(1047/5000)),
    inference(resolve, [], [refute_0_95, refute_0_102])).

cnf(refute_0_104, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_105, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) !=
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) |
     skoLx_MC1 * cos(skoAC1) != 0 * cos(1047/5000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(skoBC1)), introduced(tautology, [equality])).

cnf(refute_0_106, plain,
    (skoLx_MC1 * cos(skoAC1) != 0 * cos(1047/5000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(skoBC1)),
    inference(resolve, [], [refute_0_104, refute_0_105])).

cnf(refute_0_107, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(skoBC1)),
    inference(resolve, [], [refute_0_103, refute_0_106])).

cnf(refute_0_108, plain,
    (0 * cos(1047/5000) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(skoBC1) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000)),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_109, plain,
    (0 * cos(1047/5000) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_107, refute_0_108])).

cnf(refute_0_110, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_87, refute_0_109])).

cnf(refute_0_111, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1)), introduced(tautology, [refl])).

cnf(refute_0_112, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1)), introduced(tautology, [equality])).

cnf(refute_0_113, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1)),
    inference(resolve, [], [refute_0_111, refute_0_112])).

cnf(refute_0_114, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1)),
    inference(resolve, [], [refute_0_110, refute_0_113])).

cnf(refute_0_115, plain,
    (0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_116, plain,
    (0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_114, refute_0_115])).

cnf(refute_0_117, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) =
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000)),
    inference(resolve, [], [refute_0_83, refute_0_116])).

cnf(refute_0_118, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1))), introduced(tautology, [refl])).

cnf(refute_0_119, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) |
     skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000) |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000))),
    introduced(tautology, [equality])).

cnf(refute_0_120, plain,
    (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
     1/4 * sin(skoAC1) * cos(skoBC1) !=
     0 * cos(1047/5000) * cos(2443/10000) +
     1/4 * sin(1047/5000) * cos(2443/10000) |
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000))),
    inference(resolve, [], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000))),
    inference(resolve, [], [refute_0_117, refute_0_120])).

cnf(refute_0_122, plain, (sin(skoBC1) = sin(skoBC1)),
    introduced(tautology, [refl])).

cnf(refute_0_123, plain,
    (sin(skoBC1) != sin(skoBC1) | skoBC1 != 2443/10000 |
     sin(skoBC1) = sin(2443/10000)), introduced(tautology, [equality])).

cnf(refute_0_124, plain,
    (skoBC1 != 2443/10000 | sin(skoBC1) = sin(2443/10000)),
    inference(resolve, [], [refute_0_122, refute_0_123])).

cnf(refute_0_125, plain, (sin(skoBC1) = sin(2443/10000)),
    inference(resolve, [], [refute_0_50, refute_0_124])).

cnf(refute_0_126, plain, (sin(skoBC1) ^ 2 = sin(skoBC1) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_127, plain,
    (sin(skoBC1) ^ 2 != sin(skoBC1) ^ 2 | sin(skoBC1) != sin(2443/10000) |
     sin(skoBC1) ^ 2 = sin(2443/10000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_128, plain,
    (sin(skoBC1) != sin(2443/10000) |
     sin(skoBC1) ^ 2 = sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_126, refute_0_127])).

cnf(refute_0_129, plain, (sin(skoBC1) ^ 2 = sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_125, refute_0_128])).

cnf(refute_0_130, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2), introduced(tautology, [refl])).

cnf(refute_0_131, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 |
     sin(skoBC1) ^ 2 != sin(2443/10000) ^ 2 |
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_132, plain,
    (sin(skoBC1) ^ 2 != sin(2443/10000) ^ 2 |
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_130, refute_0_131])).

cnf(refute_0_133, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_129, refute_0_132])).

cnf(refute_0_134, plain, (sin(skoAC1) ^ 2 = sin(skoAC1) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_135, plain,
    (sin(skoAC1) ^ 2 != sin(skoAC1) ^ 2 | sin(skoAC1) != sin(1047/5000) |
     sin(skoAC1) ^ 2 = sin(1047/5000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_136, plain,
    (sin(skoAC1) != sin(1047/5000) | sin(skoAC1) ^ 2 = sin(1047/5000) ^ 2),
    inference(resolve, [], [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain, (sin(skoAC1) ^ 2 = sin(1047/5000) ^ 2),
    inference(resolve, [], [refute_0_63, refute_0_136])).

cnf(refute_0_138, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2), introduced(tautology, [refl])).

cnf(refute_0_139, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 |
     sin(skoAC1) ^ 2 != sin(1047/5000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_140, plain,
    (sin(skoAC1) ^ 2 != sin(1047/5000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2),
    inference(resolve, [], [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2),
    inference(resolve, [], [refute_0_137, refute_0_140])).

cnf(refute_0_142, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_143, plain,
    (sin(1047/5000) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_133, refute_0_143])).

cnf(refute_0_145, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2),
    introduced(tautology, [refl])).

cnf(refute_0_146, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 |
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    introduced(tautology, [equality])).

cnf(refute_0_147, plain,
    (sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_145, refute_0_146])).

cnf(refute_0_148, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 =
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2),
    inference(resolve, [], [refute_0_144, refute_0_147])).

cnf(refute_0_149, plain,
    (sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_150, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_151, plain,
    (1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2 !=
     1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2 |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_149, refute_0_150])).

cnf(refute_0_152, plain,
    (sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_148, refute_0_151])).

cnf(refute_0_153, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_154, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_155, plain,
    (sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_153, refute_0_154])).

cnf(refute_0_156, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_152, refute_0_155])).

cnf(refute_0_157, plain, (skoTMC1 = 1357143/10000),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_158, plain, (abs(skoTMC1) = abs(skoTMC1)),
    introduced(tautology, [refl])).

cnf(refute_0_159, plain,
    (abs(skoTMC1) != abs(skoTMC1) | skoTMC1 != 1357143/10000 |
     abs(skoTMC1) = abs(1357143/10000)),
    introduced(tautology, [equality])).

cnf(refute_0_160, plain,
    (skoTMC1 != 1357143/10000 | abs(skoTMC1) = abs(1357143/10000)),
    inference(resolve, [], [refute_0_158, refute_0_159])).

cnf(refute_0_161, plain, (abs(skoTMC1) = abs(1357143/10000)),
    inference(resolve, [], [refute_0_157, refute_0_160])).

cnf(refute_0_162, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [refl])).

cnf(refute_0_163, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     abs(skoTMC1) != abs(1357143/10000) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    introduced(tautology, [equality])).

cnf(refute_0_164, plain,
    (abs(skoTMC1) != abs(1357143/10000) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    inference(resolve, [], [refute_0_162, refute_0_163])).

cnf(refute_0_165, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2)),
    inference(resolve, [], [refute_0_161, refute_0_164])).

cnf(refute_0_166, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_167, plain,
    (abs(1357143/10000) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_165, refute_0_166])).

cnf(refute_0_168, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2)),
    inference(resolve, [], [refute_0_156, refute_0_167])).

cnf(refute_0_169, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1))), introduced(tautology, [refl])).

cnf(refute_0_170, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1))),
    introduced(tautology, [equality])).

cnf(refute_0_171, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1))),
    inference(resolve, [], [refute_0_169, refute_0_170])).

cnf(refute_0_172, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1))),
    inference(resolve, [], [refute_0_168, refute_0_171])).

cnf(refute_0_173, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000))),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_174, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000))),
    inference(resolve, [], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000))),
    inference(resolve, [], [refute_0_121, refute_0_174])).

cnf(refute_0_176, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1),
    introduced(tautology, [refl])).

cnf(refute_0_177, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 !=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1),
    introduced(tautology, [equality])).

cnf(refute_0_178, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1),
    inference(resolve, [], [refute_0_176, refute_0_177])).

cnf(refute_0_179, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1),
    inference(resolve, [], [refute_0_175, refute_0_178])).

cnf(refute_0_180, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5),
    inference(subst, [], [refute_0_76])).

cnf(refute_0_181, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5 |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5),
    inference(resolve, [], [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 =
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5),
    inference(resolve, [], [refute_0_49, refute_0_181])).

cnf(refute_0_183, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5 <
     abs(skoT_thetaC1) |
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1 !=
     abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5 |
     abs(skoT_thetaC1) <=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1),
    introduced(tautology, [equality])).

cnf(refute_0_184, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5 <
     abs(skoT_thetaC1) |
     abs(skoT_thetaC1) <=
     abs(skoTMC1) / sqrt(1 - sin(skoAC1) ^ 2 * sin(skoBC1) ^ 2) *
     (skoLx_MC1 * cos(skoAC1) * cos(skoBC1) +
      1/4 * sin(skoAC1) * cos(skoBC1)) - 1/50 * skoTRC1),
    inference(resolve, [], [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (abs(1357143/10000) /
     sqrt(1 - sin(1047/5000) ^ 2 * sin(2443/10000) ^ 2) *
     (0 * cos(1047/5000) * cos(2443/10000) +
      1/4 * sin(1047/5000) * cos(2443/10000)) - 1/50 * 12/5 <
     abs(skoT_thetaC1)),
    inference(resolve, [], [refute_0_184, refute_0_40])).

cnf(refute_0_186, plain,
    (-6/125 + abs(skoT_thetaC1) * -1 <
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(arithmetic, [], [refute_0_185])).

cnf(refute_0_187, plain,
    (1357143/10000 < 0 | abs(1357143/10000) = 1357143/10000),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_188, plain,
    (-6/125 + abs(skoT_thetaC1) * -1 <
     sin(1047/5000) * (cos(2443/10000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(1357143/10000) != 1357143/10000 |
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + abs(skoT_thetaC1) * -1), introduced(tautology, [equality])).

cnf(refute_0_189, plain,
    (1357143/10000 < 0 |
     -6/125 + abs(skoT_thetaC1) * -1 <
     sin(1047/5000) * (cos(2443/10000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + abs(skoT_thetaC1) * -1),
    inference(resolve, [], [refute_0_187, refute_0_188])).

cnf(refute_0_190, plain,
    (1357143/10000 < 0 |
     -6/125 + abs(skoT_thetaC1) * -1 <
     sin(1047/5000) * (cos(2443/10000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(resolve, [], [refute_0_189, refute_0_186])).

cnf(refute_0_191, plain,
    (-6/125 + abs(skoT_thetaC1) * -1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(arithmetic, [], [refute_0_190])).

cnf(refute_0_192, plain,
    (skoT_thetaC1 < 0 | abs(skoT_thetaC1) = skoT_thetaC1),
    inference(subst, [], [abs_nonnegative])).

cnf(refute_0_193, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(skoT_thetaC1) != skoT_thetaC1 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + abs(skoT_thetaC1) * -1), introduced(tautology, [equality])).

cnf(refute_0_194, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + abs(skoT_thetaC1) * -1),
    inference(resolve, [], [refute_0_192, refute_0_193])).

cnf(refute_0_195, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(resolve, [], [refute_0_194, refute_0_191])).

cnf(refute_0_196, plain,
    ((-6/125 + skoT_thetaC1 * -1) *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + skoT_thetaC1 * -1 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_197, plain,
    (skoT_thetaC1 < 0 |
     (-6/125 + skoT_thetaC1 * -1) *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(resolve, [], [refute_0_196, refute_0_195])).

cnf(refute_0_198, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     (6/125 + skoT_thetaC1) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <= -1),
    inference(arithmetic, [], [refute_0_197])).

cnf(refute_0_199, plain,
    (~ lgen(0, X_000101, sin(X_000100)) | X_000101 <= sin(X_000100)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_200, plain,
    (X_000100 < 0 | ~ lgen(0, X_000101, X_000100 - X_000100 ^ 3 / 6) |
     lgen(0, X_000101, sin(X_000100))),
    inference(subst, [], [sin_lower_bound_0_pos])).

cnf(refute_0_201, plain,
    (X_000100 < 0 | ~ lgen(0, X_000101, X_000100 - X_000100 ^ 3 / 6) |
     X_000101 <= sin(X_000100)),
    inference(resolve, [], [refute_0_200, refute_0_199])).

cnf(refute_0_202, plain,
    (X_000100 * (1 + X_000100 * (X_000100 * -1/6)) < X_000101 |
     X_000100 < 0 | X_000101 <= sin(X_000100)),
    inference(arithmetic, [], [refute_0_201])).

cnf(refute_0_203, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     0 <= sin(2443/10000)), inference(subst, [], [refute_0_202])).

cnf(refute_0_204, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <= -1),
    introduced(tautology, [assume])).

cnf(refute_0_205, plain,
    (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) < 0 |
     0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <= -1), inference(split, [], [refute_0_204])).

cnf(refute_0_206, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     0 < sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(arithmetic, [], [refute_0_205])).

cnf(refute_0_207, plain,
    (12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      -12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= -1),
    inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_208, plain,
    (sin(X_000132) < X_000133 | ~ lgen(1, sin(X_000132), X_000133)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_209, plain,
    (X_000132 < 0 |
     ~ lgen(1,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133) |
     lgen(1, sin(X_000132), X_000133)),
    inference(subst, [], [sin_upper_bound_2_pos])).

cnf(refute_0_210, plain,
    (sin(X_000132) < X_000133 | X_000132 < 0 |
     ~ lgen(1,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133)),
    inference(resolve, [], [refute_0_209, refute_0_208])).

cnf(refute_0_211, plain,
    (sin(X_000132) < X_000133 | X_000132 < 0 |
     X_000133 <=
     X_000132 *
     (1 +
      X_000132 *
      (X_000132 *
       (-1/6 +
        X_000132 *
        (X_000132 *
         (1/120 +
          X_000132 *
          (X_000132 * (-1/5040 + X_000132 * (X_000132 * 1/362880))))))))),
    inference(arithmetic, [], [refute_0_210])).

cnf(refute_0_212, plain,
    (2443/10000 < 0 |
     sin(2443/10000) <
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(subst, [], [refute_0_211])).

cnf(refute_0_213, plain,
    (12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     1 <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(subst, [], [leq_right_mul_divide_pos])).

cnf(refute_0_214, plain,
    (2443/10000 < 0 |
     sin(2443/10000) <
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(subst, [], [refute_0_211])).

cnf(refute_0_215, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <= sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     -1 / (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <=
     sin(2443/10000)), inference(split, [], [refute_0_204])).

cnf(refute_0_216, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     sin(2443/10000)), inference(arithmetic, [], [refute_0_215])).

cnf(refute_0_217, plain,
    (2443/10000 < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(resolve, [], [refute_0_216, refute_0_214])).

cnf(refute_0_218, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_217])).

cnf(refute_0_219, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     1 <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_218, refute_0_213])).

cnf(refute_0_220, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= -1 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(arithmetic, [], [refute_0_219])).

cnf(refute_0_221, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      -12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= sin(2443/10000)),
    inference(split, [], [refute_0_220])).

cnf(refute_0_222, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= sin(2443/10000)),
    inference(arithmetic, [], [refute_0_221])).

cnf(refute_0_223, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_224, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <= 0 | sin(1047/5000) = 0),
    inference(split, [], [refute_0_223])).

cnf(refute_0_225, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_224])).

cnf(refute_0_226, plain,
    (sin(1047/5000) != 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(split, [], [refute_0_5])).

cnf(refute_0_227, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_225, refute_0_226])).

cnf(refute_0_228, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= sin(2443/10000)),
    inference(resolve, [], [refute_0_222, refute_0_227])).

cnf(refute_0_229, plain,
    (2443/10000 < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <=
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880))))))))),
    inference(resolve, [], [refute_0_228, refute_0_212])).

cnf(refute_0_230, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     -1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <=
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_229])).

cnf(refute_0_231, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      -12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       -12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <= -1 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_230, refute_0_207])).

cnf(refute_0_232, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1 |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(arithmetic, [], [refute_0_231])).

cnf(refute_0_233, plain,
    (1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 <
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     sin(1047/5000) *
     -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 |
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1), inference(subst, [], [leq_right_mul_divide_neg])).

cnf(refute_0_234, plain,
    (1047/5000 < 0 |
     sin(1047/5000) <
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880))))))))),
    inference(subst, [], [refute_0_211])).

cnf(refute_0_235, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1), introduced(tautology, [assume])).

cnf(refute_0_236, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     sin(1047/5000) *
     -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)), inference(split, [], [refute_0_235])).

cnf(refute_0_237, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000) | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_236])).

cnf(refute_0_238, plain, (0 < sin(1047/5000)),
    inference(arithmetic, [], [refute_0_4])).

cnf(refute_0_239, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)),
    inference(resolve, [], [refute_0_237, refute_0_238])).

cnf(refute_0_240, plain,
    (1047/5000 < 0 |
     -1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880))))))))),
    inference(resolve, [], [refute_0_239, refute_0_234])).

cnf(refute_0_241, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     -1 /
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_240])).

cnf(refute_0_242, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     sin(1047/5000) *
     -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 |
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     -1), inference(resolve, [], [refute_0_241, refute_0_233])).

cnf(refute_0_243, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <=
     sin(1047/5000) | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_242])).

cnf(refute_0_244, plain,
    (sin(X_000072) < X_000073 | ~ lgen(1, sin(X_000072), X_000073)),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_245, plain,
    (X_000072 < 0 | ~ lgen(1, X_000072, X_000073) |
     lgen(1, sin(X_000072), X_000073)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_246, plain,
    (sin(X_000072) < X_000073 | X_000072 < 0 |
     ~ lgen(1, X_000072, X_000073)),
    inference(resolve, [], [refute_0_245, refute_0_244])).

cnf(refute_0_247, plain,
    (sin(X_000072) < X_000073 | X_000072 < 0 | X_000073 <= X_000072),
    inference(arithmetic, [], [refute_0_246])).

cnf(refute_0_248, plain,
    (1047/5000 < 0 |
     sin(1047/5000) <
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <=
     1047/5000), inference(subst, [], [refute_0_247])).

cnf(refute_0_249, plain,
    (sin(1047/5000) <
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467),
    inference(arithmetic, [], [refute_0_248])).

cnf(refute_0_250, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_243, refute_0_249])).

cnf(refute_0_251, plain,
    (1047/5000 < 0 |
     -1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) <= 0),
    inference(resolve, [], [refute_0_250, refute_0_4])).

cnf(refute_0_252, plain,
    (-1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      -157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_251])).

cnf(refute_0_253, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_232, refute_0_252])).

cnf(refute_0_254, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <= 0 |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_253])).

cnf(refute_0_255, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(1047/5000) = 0), inference(arithmetic, [], [refute_0_254])).

cnf(refute_0_256, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(resolve, [], [refute_0_255, refute_0_226])).

cnf(refute_0_257, plain,
    (0 < sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)) |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(resolve, [], [refute_0_256, refute_0_206])).

cnf(refute_0_258, plain,
    (0 < sin(2443/10000) * -1 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(split, [], [refute_0_257])).

cnf(refute_0_259, plain,
    (sin(2443/10000) < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(arithmetic, [], [refute_0_258])).

cnf(refute_0_260, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     -1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(resolve, [], [refute_0_203, refute_0_259])).

cnf(refute_0_261, plain,
    (-1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))),
    inference(arithmetic, [], [refute_0_260])).

cnf(refute_0_262, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     (6/125 + skoT_thetaC1)),
    inference(resolve, [], [refute_0_198, refute_0_261])).

cnf(refute_0_263, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     6/125 + skoT_thetaC1 <= 0), inference(split, [], [refute_0_262])).

cnf(refute_0_264, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     skoT_thetaC1 <= -6/125), inference(arithmetic, [], [refute_0_263])).

cnf(refute_0_265, plain, (-1061/10000 <= skoT_thetaC1),
    inference(canonicalize, [], [normalize_0_8])).

cnf(refute_0_266, plain, (skoT_thetaC1 <= 12667/10000),
    inference(canonicalize, [], [normalize_0_9])).

cnf(refute_0_267, plain, (skoLz_MC1 = -1/4),
    inference(canonicalize, [], [normalize_0_11])).

cnf(refute_0_268, plain, (skoT_thetaC1 < 0 | -6/125 < skoT_thetaC1),
    inference(decision, [],
              [refute_0_265, refute_0_266, refute_0_59, refute_0_50,
               refute_0_96, refute_0_157, refute_0_41, refute_0_267])).

cnf(refute_0_269, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))))),
    inference(resolve, [], [refute_0_264, refute_0_268])).

cnf(refute_0_270, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2),
    inference(resolve, [], [refute_0_39, refute_0_269])).

cnf(refute_0_271, plain,
    (skoT_thetaC1 < 0 |
     1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))),
    inference(arithmetic, [], [refute_0_270])).

cnf(refute_0_272, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <=
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_273, plain,
    (~ lgen(0, sin(X_000132), X_000133) | sin(X_000132) <= X_000133),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_274, plain,
    (X_000132 < 0 |
     ~ lgen(0,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133) |
     lgen(0, sin(X_000132), X_000133)),
    inference(subst, [], [sin_upper_bound_2_pos])).

cnf(refute_0_275, plain,
    (X_000132 < 0 |
     ~ lgen(0,
            X_000132 - X_000132 ^ 3 / 6 + X_000132 ^ 5 / 120 -
            X_000132 ^ 7 / 5040 + X_000132 ^ 9 / 362880, X_000133) |
     sin(X_000132) <= X_000133),
    inference(resolve, [], [refute_0_274, refute_0_273])).

cnf(refute_0_276, plain,
    (X_000132 < 0 |
     X_000133 <
     X_000132 *
     (1 +
      X_000132 *
      (X_000132 *
       (-1/6 +
        X_000132 *
        (X_000132 *
         (1/120 +
          X_000132 *
          (X_000132 * (-1/5040 + X_000132 * (X_000132 * 1/362880)))))))) |
     sin(X_000132) <= X_000133),
    inference(arithmetic, [], [refute_0_275])).

cnf(refute_0_277, plain,
    (2443/10000 < 0 |
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) <=
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000))),
    inference(subst, [], [refute_0_276])).

cnf(refute_0_278, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <=
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_279, plain,
    (2443/10000 < 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) <=
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)))),
    inference(subst, [], [refute_0_276])).

cnf(refute_0_280, plain,
    (1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    introduced(tautology, [assume])).

cnf(refute_0_281, plain,
    (1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     sin(2443/10000) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(split, [], [refute_0_280])).

cnf(refute_0_282, plain,
    (1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     sin(2443/10000) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_281, refute_0_14])).

cnf(refute_0_283, plain,
    (2443/10000 < 0 |
     1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_279, refute_0_282])).

cnf(refute_0_284, plain,
    (1 / (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_283])).

cnf(refute_0_285, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_278, refute_0_284])).

cnf(refute_0_286, plain,
    (1 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(2443/10000) * (sin(1047/5000) * sin(1047/5000)) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_285])).

cnf(refute_0_287, plain,
    (1 <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_286, refute_0_14])).

cnf(refute_0_288, plain,
    (1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) < sin(2443/10000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(split, [], [refute_0_287])).

cnf(refute_0_289, plain, (sin(1047/5000) != 0),
    inference(split, [], [refute_0_14])).

cnf(refute_0_290, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_225, refute_0_289])).

cnf(refute_0_291, plain,
    (1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) < sin(2443/10000) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_288, refute_0_290])).

cnf(refute_0_292, plain,
    (2443/10000 < 0 |
     1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 + 2443/10000 * (2443/10000 * 1/362880)))))))) |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_277, refute_0_291])).

cnf(refute_0_293, plain,
    (1 /
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_292])).

cnf(refute_0_294, plain,
    (1 <
     12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       12538912456006985475956782657218819948349/
       51840000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_272, refute_0_293])).

cnf(refute_0_295, plain,
    (1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(arithmetic, [], [refute_0_294])).

cnf(refute_0_296, plain,
    (~ lgen(0, sin(X_000072), X_000073) | sin(X_000072) <= X_000073),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_297, plain,
    (X_000072 < 0 | ~ lgen(0, X_000072, X_000073) |
     lgen(0, sin(X_000072), X_000073)),
    inference(subst, [], [sin_upper_bound_0_pos])).

cnf(refute_0_298, plain,
    (X_000072 < 0 | ~ lgen(0, X_000072, X_000073) |
     sin(X_000072) <= X_000073),
    inference(resolve, [], [refute_0_297, refute_0_296])).

cnf(refute_0_299, plain,
    (X_000073 < X_000072 | X_000072 < 0 | sin(X_000072) <= X_000073),
    inference(arithmetic, [], [refute_0_298])).

cnf(refute_0_300, plain,
    (1047/5000 < 0 |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     1047/5000 |
     sin(1047/5000) <=
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467),
    inference(subst, [], [refute_0_299])).

cnf(refute_0_301, plain,
    (1 <
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 <=
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_302, plain,
    (1047/5000 < 0 |
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880)))))))) |
     sin(1047/5000) <=
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(subst, [], [refute_0_276])).

cnf(refute_0_303, plain,
    (1 <
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), introduced(tautology, [assume])).

cnf(refute_0_304, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     sin(1047/5000) *
     157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(split, [], [refute_0_303])).

cnf(refute_0_305, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1 | sin(1047/5000) <= 0), inference(arithmetic, [], [refute_0_304])).

cnf(refute_0_306, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_305, refute_0_238])).

cnf(refute_0_307, plain,
    (1047/5000 < 0 |
     1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 + 1047/5000 * (1047/5000 * 1/362880)))))))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_302, refute_0_306])).

cnf(refute_0_308, plain,
    (1 /
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(arithmetic, [], [refute_0_307])).

cnf(refute_0_309, plain,
    (1 <
     1818889133334824227556980993206135207/
     8750000000000000000000000000000000000 *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     157224325579407132478971229953446022324529409796000092373769800136381459027825801/
     2687385600000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_301, refute_0_308])).

cnf(refute_0_310, plain,
    (1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1 | sin(1047/5000) <= 0), inference(arithmetic, [], [refute_0_309])).

cnf(refute_0_311, plain,
    (1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     sin(1047/5000) |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_310, refute_0_238])).

cnf(refute_0_312, plain,
    (1047/5000 < 0 |
     1119744000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000/
     13617791299632384533999257259519435982415723080791165932336798925414084386634564185662091328954746502841108726146467 <
     1047/5000 |
     sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(resolve, [], [refute_0_300, refute_0_311])).

cnf(refute_0_313, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      157224325579407132478971229953446022324529409796000092373769800136381459027825801/
      2687385600000000000000000000000000000000000000000000000000000000000000000000000000) <=
     1), inference(arithmetic, [], [refute_0_312])).

cnf(refute_0_314, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      12538912456006985475956782657218819948349/
      51840000000000000000000000000000000000000) <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_313, refute_0_295])).

cnf(refute_0_315, plain,
    (12538912456006985475956782657218819948349/
     51840000000000000000000000000000000000000 <= 0 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_314])).

cnf(refute_0_316, plain,
    (sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1 |
     sin(1047/5000) = 0), inference(arithmetic, [], [refute_0_315])).

cnf(refute_0_317, plain,
    (sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) <= 1),
    inference(resolve, [], [refute_0_316, refute_0_289])).

cnf(refute_0_318, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))),
    inference(resolve, [], [refute_0_317, refute_0_271])).

cnf(refute_0_319, plain,
    (skoT_thetaC1 < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (6/125 + skoT_thetaC1) | 6/125 + skoT_thetaC1 <= 0),
    inference(resolve, [], [refute_0_34, refute_0_318])).

cnf(refute_0_320, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     skoT_thetaC1 <= -6/125), inference(arithmetic, [], [refute_0_319])).

cnf(refute_0_321, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))))),
    inference(resolve, [], [refute_0_320, refute_0_268])).

cnf(refute_0_322, plain,
    (skoT_thetaC1 < 0 |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))),
    inference(split, [], [refute_0_321])).

cnf(refute_0_323, plain,
    (skoT_thetaC1 < 0 |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_322])).

cnf(refute_0_324, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     cos(2443/10000)), inference(subst, [], [refute_0_20])).

cnf(refute_0_325, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_326, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     sin(1047/5000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_327, plain,
    (-6/125 + skoT_thetaC1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_328, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) <=
     sin(1047/5000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_329, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_330, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_331, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) <=
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_332, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) <=
     sin(2443/10000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_333, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (6/125 + skoT_thetaC1 * -1) | 6/125 + skoT_thetaC1 * -1 <= 0 |
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) <=
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1)),
    inference(subst, [], [leq_right_divide_mul_pos])).

cnf(refute_0_334, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1)), inference(subst, [], [refute_0_38])).

cnf(refute_0_335, plain,
    (0 <= skoT_thetaC1 | abs(skoT_thetaC1) = -skoT_thetaC1),
    inference(subst, [], [abs_negative])).

cnf(refute_0_336, plain,
    (-6/125 + -skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(skoT_thetaC1) != -skoT_thetaC1 |
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + abs(skoT_thetaC1) * -1), introduced(tautology, [equality])).

cnf(refute_0_337, plain,
    (-6/125 + -skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1 |
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + abs(skoT_thetaC1) * -1),
    inference(resolve, [], [refute_0_335, refute_0_336])).

cnf(refute_0_338, plain,
    (-6/125 + -skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_337, refute_0_186])).

cnf(refute_0_339, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1), inference(arithmetic, [], [refute_0_338])).

cnf(refute_0_340, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     abs(1357143/10000) != 1357143/10000 |
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + skoT_thetaC1), introduced(tautology, [equality])).

cnf(refute_0_341, plain,
    (1357143/10000 < 0 |
     -6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     sin(1047/5000) * (cos(2443/10000) * (abs(1357143/10000) * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + skoT_thetaC1),
    inference(resolve, [], [refute_0_187, refute_0_340])).

cnf(refute_0_342, plain,
    (1357143/10000 < 0 |
     -6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * (1357143/10000 * -1/4)) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_341, refute_0_339])).

cnf(refute_0_343, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1), inference(arithmetic, [], [refute_0_342])).

cnf(refute_0_344, plain,
    ((-6/125 + skoT_thetaC1) *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) /
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     -6/125 + skoT_thetaC1 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(subst, [], [leq_left_divide_mul_pos])).

cnf(refute_0_345, plain,
    ((-6/125 + skoT_thetaC1) *
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <= skoT_thetaC1 |
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) <=
     0), inference(resolve, [], [refute_0_344, refute_0_343])).

cnf(refute_0_346, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     (6/125 + skoT_thetaC1 * -1) | 0 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) <= -1),
    inference(arithmetic, [], [refute_0_345])).

cnf(refute_0_347, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) *
     (6/125 + skoT_thetaC1 * -1) | 0 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_346, refute_0_261])).

cnf(refute_0_348, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1 | 6/125 + skoT_thetaC1 * -1 <= 0),
    inference(split, [], [refute_0_347])).

cnf(refute_0_349, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1 | 6/125 <= skoT_thetaC1),
    inference(arithmetic, [], [refute_0_348])).

cnf(refute_0_350, plain, (skoT_thetaC1 < 6/125 | 0 <= skoT_thetaC1),
    inference(decision, [],
              [refute_0_267, refute_0_41, refute_0_157, refute_0_96,
               refute_0_50, refute_0_59, refute_0_266, refute_0_265])).

cnf(refute_0_351, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     sqrt(1 +
          sin(2443/10000) *
          (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) |
     0 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_349, refute_0_350])).

cnf(refute_0_352, plain,
    (1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1))) < 0 |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     1/2 +
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1)))) * 1/2 |
     0 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_334, refute_0_351])).

cnf(refute_0_353, plain,
    (1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     0 <= skoT_thetaC1), inference(arithmetic, [], [refute_0_352])).

cnf(refute_0_354, plain,
    (sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <
     0 |
     -6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) < 0 |
     skoT_thetaC1 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))),
    inference(split, [], [refute_0_321])).

cnf(refute_0_355, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2)))),
    inference(arithmetic, [], [refute_0_354])).

cnf(refute_0_356, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 <=
     cos(2443/10000)), inference(subst, [], [refute_0_20])).

cnf(refute_0_357, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000) |
     0 <= cos(2443/10000) * -1357143/40000 |
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_358, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) |
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) <=
     sin(1047/5000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_359, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1), introduced(tautology, [assume])).

cnf(refute_0_360, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) |
     0 <= cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1), inference(split, [], [refute_0_359])).

cnf(refute_0_361, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1 | cos(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_360])).

cnf(refute_0_362, plain,
    (X_000015 < cos(X_000014) | ~ lgen(1, X_000015, cos(X_000014))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_363, plain,
    (~ lgen(1, X_000015, 1 - X_000014 ^ 2 / 2) |
     lgen(1, X_000015, cos(X_000014))),
    inference(subst, [], [cos_lower_bound_0])).

cnf(refute_0_364, plain,
    (X_000015 < cos(X_000014) | ~ lgen(1, X_000015, 1 - X_000014 ^ 2 / 2)),
    inference(resolve, [], [refute_0_363, refute_0_362])).

cnf(refute_0_365, plain,
    (X_000015 < cos(X_000014) |
     1 + X_000014 * (X_000014 * -1/2) <= X_000015),
    inference(arithmetic, [], [refute_0_364])).

cnf(refute_0_366, plain,
    (0 < cos(2443/10000) | 1 + 2443/10000 * (2443/10000 * -1/2) <= 0),
    inference(subst, [], [refute_0_365])).

cnf(refute_0_367, plain, (0 < cos(2443/10000)),
    inference(arithmetic, [], [refute_0_366])).

cnf(refute_0_368, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(resolve, [], [refute_0_361, refute_0_367])).

cnf(refute_0_369, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(resolve, [], [refute_0_358, refute_0_368])).

cnf(refute_0_370, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(arithmetic, [], [refute_0_369])).

cnf(refute_0_371, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000) |
     0 <= cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(resolve, [], [refute_0_357, refute_0_370])).

cnf(refute_0_372, plain,
    (cos(2443/10000) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1 | cos(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_371])).

cnf(refute_0_373, plain,
    (cos(2443/10000) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(resolve, [], [refute_0_372, refute_0_367])).

cnf(refute_0_374, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(resolve, [], [refute_0_356, refute_0_373])).

cnf(refute_0_375, plain,
    (1307916368539235571037117871615276716433805369757135585645641/
     192500000000000000000000000000000000000000000000000000000000 <
     skoT_thetaC1 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(arithmetic, [], [refute_0_374])).

cnf(refute_0_376, plain,
    (skoT_thetaC1 <=
     1307916368539235571037117871615276716433805369757135585645641/
     192500000000000000000000000000000000000000000000000000000000),
    inference(decision, [],
              [refute_0_267, refute_0_41, refute_0_157, refute_0_96,
               refute_0_50, refute_0_59, refute_0_266, refute_0_265])).

cnf(refute_0_377, plain,
    (sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 * -1),
    inference(resolve, [], [refute_0_376, refute_0_375])).

cnf(refute_0_378, plain,
    (skoT_thetaC1 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2)))),
    inference(resolve, [], [refute_0_377, refute_0_355])).

cnf(refute_0_379, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2))) |
     0 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_317, refute_0_353])).

cnf(refute_0_380, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (6/125 + skoT_thetaC1 * -1) | 0 <= skoT_thetaC1 |
     6/125 + skoT_thetaC1 * -1 <= 0),
    inference(resolve, [], [refute_0_333, refute_0_379])).

cnf(refute_0_381, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <= skoT_thetaC1 | 6/125 <= skoT_thetaC1),
    inference(arithmetic, [], [refute_0_380])).

cnf(refute_0_382, plain, (0 <= sin(2443/10000)),
    inference(arithmetic, [], [refute_0_203])).

cnf(refute_0_383, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_384, plain,
    (0 < sin(2443/10000) * (-3/125 + skoT_thetaC1 * -1/2) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_383])).

cnf(refute_0_385, plain,
    (sin(2443/10000) < 0 / (-3/125 + skoT_thetaC1 * -1/2) |
     0 <= -3/125 + skoT_thetaC1 * -1/2 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_384])).

cnf(refute_0_386, plain,
    (sin(2443/10000) < 0 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     skoT_thetaC1 <= -6/125), inference(arithmetic, [], [refute_0_385])).

cnf(refute_0_387, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     skoT_thetaC1 <= -6/125),
    inference(resolve, [], [refute_0_382, refute_0_386])).

cnf(refute_0_388, plain,
    (0 < 0 | -3/125 + skoT_thetaC1 * -1/2 < 0 |
     0 < -3/125 + skoT_thetaC1 * -1/2 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_384])).

cnf(refute_0_389, plain,
    (skoT_thetaC1 < -6/125 | -6/125 < skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_388])).

cnf(refute_0_390, plain,
    (skoT_thetaC1 < 0 | 6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(decision, [],
              [refute_0_387, refute_0_389, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267])).

cnf(refute_0_391, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_381, refute_0_390])).

cnf(refute_0_392, plain,
    (skoT_thetaC1 < 6/125 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(decision, [],
              [refute_0_387, refute_0_389, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267])).

cnf(refute_0_393, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_391, refute_0_392])).

cnf(refute_0_394, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_393])).

cnf(refute_0_395, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), inference(arithmetic, [], [refute_0_394])).

cnf(refute_0_396, plain,
    (sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <
     0 |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_393])).

cnf(refute_0_397, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_396])).

cnf(refute_0_398, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     -962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     -962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 <=
     cos(2443/10000)), inference(subst, [], [refute_0_20])).

cnf(refute_0_399, plain,
    (-6/125 + skoT_thetaC1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000) |
     0 <= cos(2443/10000) * -1357143/40000 |
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) <=
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000),
    inference(subst, [], [leq_right_divide_mul_neg])).

cnf(refute_0_400, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) |
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) <=
     sin(1047/5000)), inference(subst, [], [refute_0_26])).

cnf(refute_0_401, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1), introduced(tautology, [assume])).

cnf(refute_0_402, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) |
     0 <= cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1), inference(split, [], [refute_0_401])).

cnf(refute_0_403, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 | cos(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_402])).

cnf(refute_0_404, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1),
    inference(resolve, [], [refute_0_403, refute_0_367])).

cnf(refute_0_405, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1),
    inference(resolve, [], [refute_0_400, refute_0_404])).

cnf(refute_0_406, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1) / (cos(2443/10000) * -1357143/40000) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1), inference(arithmetic, [], [refute_0_405])).

cnf(refute_0_407, plain,
    (-6/125 + skoT_thetaC1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000) |
     0 <= cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1),
    inference(resolve, [], [refute_0_399, refute_0_406])).

cnf(refute_0_408, plain,
    (cos(2443/10000) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     -962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 | cos(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_407])).

cnf(refute_0_409, plain,
    (cos(2443/10000) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     -962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1),
    inference(resolve, [], [refute_0_408, refute_0_367])).

cnf(refute_0_410, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     15400000000000000000000000000000000000000000000000/
     2262784933824628820735492016828518470041200232493797 +
     skoT_thetaC1 *
     -962500000000000000000000000000000000000000000000000/
     6788354801473886462206476050485555410123600697481391 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1),
    inference(resolve, [], [refute_0_398, refute_0_409])).

cnf(refute_0_411, plain,
    (skoT_thetaC1 <
     -1307916368539235571037117871615276716433805369757135585645641/
     192500000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1), inference(arithmetic, [], [refute_0_410])).

cnf(refute_0_412, plain,
    (-1307916368539235571037117871615276716433805369757135585645641/
     192500000000000000000000000000000000000000000000000000000000 <=
     skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(decision, [],
              [refute_0_387, refute_0_389, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267])).

cnf(refute_0_413, plain,
    (sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_412, refute_0_411])).

cnf(refute_0_414, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_413, refute_0_397])).

cnf(refute_0_415, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_416, plain,
    (0 < sin(2443/10000) * (-3/125 + skoT_thetaC1 * 1/2) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(split, [], [refute_0_415])).

cnf(refute_0_417, plain,
    (sin(2443/10000) < 0 / (-3/125 + skoT_thetaC1 * 1/2) |
     0 <= -3/125 + skoT_thetaC1 * 1/2 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(split, [], [refute_0_416])).

cnf(refute_0_418, plain,
    (sin(2443/10000) < 0 | 6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_417])).

cnf(refute_0_419, plain,
    (sin(2443/10000) < 0 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_418, refute_0_392])).

cnf(refute_0_420, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_203, refute_0_419])).

cnf(refute_0_421, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_420])).

cnf(refute_0_422, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_421, refute_0_414])).

cnf(refute_0_423, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_395, refute_0_422])).

cnf(refute_0_424, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_332, refute_0_423])).

cnf(refute_0_425, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_424])).

cnf(refute_0_426, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_331, refute_0_425])).

cnf(refute_0_427, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), inference(arithmetic, [], [refute_0_426])).

cnf(refute_0_428, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_427, refute_0_422])).

cnf(refute_0_429, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_428])).

cnf(refute_0_430, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_429])).

cnf(refute_0_431, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) < 0 |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) < 0 |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_428])).

cnf(refute_0_432, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_431])).

cnf(refute_0_433, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_413, refute_0_432])).

cnf(refute_0_434, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_435, plain,
    (0 <
     -137928037016074197461091406895432752238788286029099/
     23760000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(split, [], [refute_0_434])).

cnf(refute_0_436, plain,
    (6/125 < skoT_thetaC1 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    inference(arithmetic, [], [refute_0_435])).

cnf(refute_0_437, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) * (-3/125 + skoT_thetaC1 * 1/2) <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_421])).

cnf(refute_0_438, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) * (-3/125 + skoT_thetaC1 * 1/2) <= 0),
    inference(resolve, [], [refute_0_437, refute_0_289])).

cnf(refute_0_439, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     -3/125 + skoT_thetaC1 * 1/2 <= 0 |
     sin(2443/10000) <= 0 / (-3/125 + skoT_thetaC1 * 1/2)),
    inference(split, [], [refute_0_438])).

cnf(refute_0_440, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(2443/10000) <= 0 | skoT_thetaC1 <= 6/125),
    inference(arithmetic, [], [refute_0_439])).

cnf(refute_0_441, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     skoT_thetaC1 <= 6/125),
    inference(resolve, [], [refute_0_440, refute_0_9])).

cnf(refute_0_442, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_441, refute_0_436])).

cnf(refute_0_443, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_442, refute_0_433])).

cnf(refute_0_444, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_430, refute_0_443])).

cnf(refute_0_445, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_330, refute_0_444])).

cnf(refute_0_446, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_445])).

cnf(refute_0_447, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_329, refute_0_446])).

cnf(refute_0_448, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_447])).

cnf(refute_0_449, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_448, refute_0_443])).

cnf(refute_0_450, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_449])).

cnf(refute_0_451, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_450])).

cnf(refute_0_452, plain,
    (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <
     cos(2443/10000) |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(subst, [], [refute_0_365])).

cnf(refute_0_453, plain,
    (cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_454, plain,
    (X_000168 < sin(X_000167) | ~ lgen(1, X_000168, sin(X_000167))),
    inference(subst, [], [lgen_less_neg])).

cnf(refute_0_455, plain,
    (X_000167 < 0 |
     ~ lgen(1, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800) | lgen(1, X_000168, sin(X_000167))),
    inference(subst, [], [sin_lower_bound_2_pos])).

cnf(refute_0_456, plain,
    (X_000167 < 0 | X_000168 < sin(X_000167) |
     ~ lgen(1, X_000168,
            X_000167 - X_000167 ^ 3 / 6 + X_000167 ^ 5 / 120 -
            X_000167 ^ 7 / 5040 + X_000167 ^ 9 / 362880 -
            X_000167 ^ 11 / 39916800)),
    inference(resolve, [], [refute_0_455, refute_0_454])).

cnf(refute_0_457, plain,
    (X_000167 < 0 | X_000168 < sin(X_000167) |
     X_000167 *
     (1 +
      X_000167 *
      (X_000167 *
       (-1/6 +
        X_000167 *
        (X_000167 *
         (1/120 +
          X_000167 *
          (X_000167 *
           (-1/5040 +
            X_000167 *
            (X_000167 *
             (1/362880 + X_000167 * (X_000167 * -1/39916800)))))))))) <=
     X_000168), inference(arithmetic, [], [refute_0_456])).

cnf(refute_0_458, plain,
    (1047/5000 < 0 |
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(subst, [], [refute_0_457])).

cnf(refute_0_459, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000), introduced(tautology, [assume])).

cnf(refute_0_460, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_459])).

cnf(refute_0_461, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     6/125 <= skoT_thetaC1 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(arithmetic, [], [refute_0_460])).

cnf(refute_0_462, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_461, refute_0_392])).

cnf(refute_0_463, plain,
    (1047/5000 < 0 |
     cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_462, refute_0_458])).

cnf(refute_0_464, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_463])).

cnf(refute_0_465, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(resolve, [], [refute_0_464, refute_0_453])).

cnf(refute_0_466, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_465])).

cnf(refute_0_467, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_466, refute_0_392])).

cnf(refute_0_468, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0 |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_467, refute_0_452])).

cnf(refute_0_469, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 <=
     skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_468])).

cnf(refute_0_470, plain,
    (skoT_thetaC1 <
     3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(decision, [],
              [refute_0_267, refute_0_41, refute_0_157, refute_0_96,
               refute_0_50, refute_0_59, refute_0_266, refute_0_265,
               refute_0_389, refute_0_387])).

cnf(refute_0_471, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_469, refute_0_470])).

cnf(refute_0_472, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_451, refute_0_471])).

cnf(refute_0_473, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_328, refute_0_472])).

cnf(refute_0_474, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_473])).

cnf(refute_0_475, plain,
    (-6/125 + skoT_thetaC1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_327, refute_0_474])).

cnf(refute_0_476, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_475])).

cnf(refute_0_477, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_476, refute_0_471])).

cnf(refute_0_478, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(split, [], [refute_0_477])).

cnf(refute_0_479, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_478])).

cnf(refute_0_480, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_479, refute_0_392])).

cnf(refute_0_481, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_326, refute_0_480])).

cnf(refute_0_482, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_481])).

cnf(refute_0_483, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_325, refute_0_482])).

cnf(refute_0_484, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_483])).

cnf(refute_0_485, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_484, refute_0_392])).

cnf(refute_0_486, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_324, refute_0_485])).

cnf(refute_0_487, plain,
    (skoT_thetaC1 <
     -5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(arithmetic, [], [refute_0_486])).

cnf(refute_0_488, plain,
    (-5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 <=
     skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(decision, [],
              [refute_0_267, refute_0_41, refute_0_157, refute_0_96,
               refute_0_50, refute_0_59, refute_0_266, refute_0_265,
               refute_0_389, refute_0_387])).

cnf(refute_0_489, plain,
    (sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) <= 0),
    inference(resolve, [], [refute_0_488, refute_0_487])).

cnf(refute_0_490, plain,
    (skoT_thetaC1 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2)))),
    inference(resolve, [], [refute_0_489, refute_0_378])).

cnf(refute_0_491, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     1 <
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * sin(1047/5000))) |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))),
    inference(resolve, [], [refute_0_353, refute_0_490])).

cnf(refute_0_492, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) /
     (6/125 + skoT_thetaC1 * -1) <
     1 +
     sin(2443/10000) *
     (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))),
    inference(resolve, [], [refute_0_317, refute_0_491])).

cnf(refute_0_493, plain,
    (sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     (1 +
      sin(2443/10000) *
      (sin(2443/10000) * (sin(1047/5000) * (sin(1047/5000) * -1/2)))) *
     (6/125 + skoT_thetaC1 * -1) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     6/125 + skoT_thetaC1 * -1 <= 0),
    inference(resolve, [], [refute_0_333, refute_0_492])).

cnf(refute_0_494, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     6/125 <= skoT_thetaC1), inference(arithmetic, [], [refute_0_493])).

cnf(refute_0_495, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_496, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) * (3/125 + skoT_thetaC1 * 1/2) <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_495])).

cnf(refute_0_497, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) * (3/125 + skoT_thetaC1 * 1/2) <= 0),
    inference(resolve, [], [refute_0_496, refute_0_289])).

cnf(refute_0_498, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     3/125 + skoT_thetaC1 * 1/2 <= 0 |
     sin(2443/10000) <= 0 / (3/125 + skoT_thetaC1 * 1/2)),
    inference(split, [], [refute_0_497])).

cnf(refute_0_499, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <= 0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_498])).

cnf(refute_0_500, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     skoT_thetaC1 <= -6/125),
    inference(resolve, [], [refute_0_499, refute_0_9])).

cnf(refute_0_501, plain,
    (0 < sin(1047/5000) * sin(1047/5000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_502, plain,
    (0 < 0 | sin(1047/5000) < 0 | 0 < sin(1047/5000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_501])).

cnf(refute_0_503, plain,
    (sin(1047/5000) < 0 | 0 < sin(1047/5000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_502])).

cnf(refute_0_504, plain,
    (1047/5000 < 0 | 1047/5000 * (1 + 1047/5000 * (1047/5000 * -1/6)) < 0 |
     0 <= sin(1047/5000)), inference(subst, [], [refute_0_202])).

cnf(refute_0_505, plain, (0 <= sin(1047/5000)),
    inference(arithmetic, [], [refute_0_504])).

cnf(refute_0_506, plain,
    (0 < sin(1047/5000) | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_505, refute_0_503])).

cnf(refute_0_507, plain,
    (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)) <=
     0 |
     sin(2443/10000) <=
     0 /
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))),
    inference(split, [], [refute_0_489])).

cnf(refute_0_508, plain,
    (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)) <=
     0 | sin(2443/10000) <= 0), inference(arithmetic, [], [refute_0_507])).

cnf(refute_0_509, plain,
    (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)) <=
     0), inference(resolve, [], [refute_0_508, refute_0_9])).

cnf(refute_0_510, plain,
    (-3/125 + skoT_thetaC1 * -1/2 <= 0 | sin(1047/5000) = 0),
    inference(split, [], [refute_0_509])).

cnf(refute_0_511, plain, (-6/125 <= skoT_thetaC1 | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_510])).

cnf(refute_0_512, plain,
    (0 < sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) |
     sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_513, plain,
    (0 < 0 | -3/125 + skoT_thetaC1 * -1/2 < 0 |
     0 < -3/125 + skoT_thetaC1 * -1/2 |
     sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0),
    inference(split, [], [refute_0_512])).

cnf(refute_0_514, plain,
    (skoT_thetaC1 < -6/125 | -6/125 < skoT_thetaC1 |
     sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0),
    inference(arithmetic, [], [refute_0_513])).

cnf(refute_0_515, plain,
    (sin(1047/5000) < 0 / (-3/125 + skoT_thetaC1 * -1/2) |
     0 <= -3/125 + skoT_thetaC1 * -1/2 |
     sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0),
    inference(split, [], [refute_0_512])).

cnf(refute_0_516, plain,
    (sin(1047/5000) < 0 |
     sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0 |
     skoT_thetaC1 <= -6/125), inference(arithmetic, [], [refute_0_515])).

cnf(refute_0_517, plain,
    (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0 |
     skoT_thetaC1 <= -6/125),
    inference(resolve, [], [refute_0_505, refute_0_516])).

cnf(refute_0_518, plain,
    (skoT_thetaC1 < -6/125 |
     sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_517, refute_0_514])).

cnf(refute_0_519, plain,
    (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0 |
     sin(1047/5000) = 0),
    inference(resolve, [], [refute_0_511, refute_0_518])).

cnf(refute_0_520, plain,
    (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_519, refute_0_289])).

cnf(refute_0_521, plain,
    (-3/125 + skoT_thetaC1 * -1/2 <= 0 |
     sin(1047/5000) <= 0 / (-3/125 + skoT_thetaC1 * -1/2)),
    inference(split, [], [refute_0_520])).

cnf(refute_0_522, plain, (-6/125 <= skoT_thetaC1 | sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_521])).

cnf(refute_0_523, plain,
    (-6/125 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_522, refute_0_506])).

cnf(refute_0_524, plain,
    (skoT_thetaC1 < 6/125 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_500, refute_0_490, refute_0_523, refute_0_267,
               refute_0_41, refute_0_157, refute_0_96, refute_0_50,
               refute_0_59, refute_0_266, refute_0_265])).

cnf(refute_0_525, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_494, refute_0_524])).

cnf(refute_0_526, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_525])).

cnf(refute_0_527, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), inference(arithmetic, [], [refute_0_526])).

cnf(refute_0_528, plain,
    (sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <
     0 |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_525])).

cnf(refute_0_529, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_528])).

cnf(refute_0_530, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     -1307916368539235571037117871615276716433805369757135585645641/
     192500000000000000000000000000000000000000000000000000000000 <=
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_265, refute_0_266, refute_0_59, refute_0_50,
               refute_0_96, refute_0_157, refute_0_41, refute_0_267,
               refute_0_523, refute_0_490, refute_0_500])).

cnf(refute_0_531, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_530, refute_0_411])).

cnf(refute_0_532, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_531, refute_0_529])).

cnf(refute_0_533, plain,
    (sin(2443/10000) < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_418, refute_0_524])).

cnf(refute_0_534, plain,
    (2443/10000 < 0 |
     2443/10000 * (1 + 2443/10000 * (2443/10000 * -1/6)) < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_203, refute_0_533])).

cnf(refute_0_535, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_534])).

cnf(refute_0_536, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_535, refute_0_532])).

cnf(refute_0_537, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_527, refute_0_536])).

cnf(refute_0_538, plain,
    (2443/10000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_332, refute_0_537])).

cnf(refute_0_539, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_538])).

cnf(refute_0_540, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_331, refute_0_539])).

cnf(refute_0_541, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), inference(arithmetic, [], [refute_0_540])).

cnf(refute_0_542, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_541, refute_0_536])).

cnf(refute_0_543, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_542])).

cnf(refute_0_544, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_543])).

cnf(refute_0_545, plain,
    (sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) < 0 |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) < 0 |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_542])).

cnf(refute_0_546, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_545])).

cnf(refute_0_547, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_531, refute_0_546])).

cnf(refute_0_548, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) * (-3/125 + skoT_thetaC1 * 1/2) <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_535])).

cnf(refute_0_549, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) * (-3/125 + skoT_thetaC1 * 1/2) <= 0),
    inference(resolve, [], [refute_0_548, refute_0_289])).

cnf(refute_0_550, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     -3/125 + skoT_thetaC1 * 1/2 <= 0 |
     sin(2443/10000) <= 0 / (-3/125 + skoT_thetaC1 * 1/2)),
    inference(split, [], [refute_0_549])).

cnf(refute_0_551, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 | sin(2443/10000) <= 0 |
     skoT_thetaC1 <= 6/125), inference(arithmetic, [], [refute_0_550])).

cnf(refute_0_552, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= 6/125),
    inference(resolve, [], [refute_0_551, refute_0_9])).

cnf(refute_0_553, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_552, refute_0_436])).

cnf(refute_0_554, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_553, refute_0_547])).

cnf(refute_0_555, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_544, refute_0_554])).

cnf(refute_0_556, plain,
    (2443/10000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_330, refute_0_555])).

cnf(refute_0_557, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_556])).

cnf(refute_0_558, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_329, refute_0_557])).

cnf(refute_0_559, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_558])).

cnf(refute_0_560, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_559, refute_0_554])).

cnf(refute_0_561, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_560])).

cnf(refute_0_562, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_561])).

cnf(refute_0_563, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_461, refute_0_524])).

cnf(refute_0_564, plain,
    (1047/5000 < 0 |
     cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_563, refute_0_458])).

cnf(refute_0_565, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_564])).

cnf(refute_0_566, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(resolve, [], [refute_0_565, refute_0_453])).

cnf(refute_0_567, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     6/125 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_566])).

cnf(refute_0_568, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_567, refute_0_524])).

cnf(refute_0_569, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_568, refute_0_452])).

cnf(refute_0_570, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 <=
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_569])).

cnf(refute_0_571, plain,
    (skoT_thetaC1 <
     3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_265, refute_0_266, refute_0_59, refute_0_50,
               refute_0_96, refute_0_157, refute_0_41, refute_0_267,
               refute_0_523, refute_0_490, refute_0_500])).

cnf(refute_0_572, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_570, refute_0_571])).

cnf(refute_0_573, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_562, refute_0_572])).

cnf(refute_0_574, plain,
    (1047/5000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_328, refute_0_573])).

cnf(refute_0_575, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_574])).

cnf(refute_0_576, plain,
    (-6/125 + skoT_thetaC1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_327, refute_0_575])).

cnf(refute_0_577, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_576])).

cnf(refute_0_578, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_577, refute_0_572])).

cnf(refute_0_579, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_578])).

cnf(refute_0_580, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     6/125 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_579])).

cnf(refute_0_581, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_580, refute_0_524])).

cnf(refute_0_582, plain,
    (1047/5000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_326, refute_0_581])).

cnf(refute_0_583, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_582])).

cnf(refute_0_584, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_325, refute_0_583])).

cnf(refute_0_585, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     6/125 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_584])).

cnf(refute_0_586, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_585, refute_0_524])).

cnf(refute_0_587, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_324, refute_0_586])).

cnf(refute_0_588, plain,
    (skoT_thetaC1 <
     -5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_587])).

cnf(refute_0_589, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     -5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 <=
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_265, refute_0_266, refute_0_59, refute_0_50,
               refute_0_96, refute_0_157, refute_0_41, refute_0_267,
               refute_0_523, refute_0_490, refute_0_500])).

cnf(refute_0_590, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_589, refute_0_588])).

cnf(refute_0_591, plain,
    (skoT_thetaC1 < 0 |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_323, refute_0_590])).

cnf(refute_0_592, plain,
    (2443/10000 < 0 | skoT_thetaC1 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_33, refute_0_591])).

cnf(refute_0_593, plain,
    (skoT_thetaC1 < 0 |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_592])).

cnf(refute_0_594, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_32, refute_0_593])).

cnf(refute_0_595, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_594])).

cnf(refute_0_596, plain,
    (skoT_thetaC1 < 0 |
     -6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_595, refute_0_590])).

cnf(refute_0_597, plain,
    (skoT_thetaC1 < 0 |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_596])).

cnf(refute_0_598, plain,
    (skoT_thetaC1 < 0 |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_597])).

cnf(refute_0_599, plain,
    (~ lgen(0, X_000029, cos(X_000028)) | X_000029 <= cos(X_000028)),
    inference(subst, [], [lgen_le_neg])).

cnf(refute_0_600, plain,
    (~ lgen(0, X_000029,
            1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24 - X_000028 ^ 6 / 720 +
            X_000028 ^ 8 / 40320 - X_000028 ^ 10 / 3628800) |
     lgen(0, X_000029, cos(X_000028))),
    inference(subst, [], [cos_lower_bound_2])).

cnf(refute_0_601, plain,
    (~ lgen(0, X_000029,
            1 - X_000028 ^ 2 / 2 + X_000028 ^ 4 / 24 - X_000028 ^ 6 / 720 +
            X_000028 ^ 8 / 40320 - X_000028 ^ 10 / 3628800) |
     X_000029 <= cos(X_000028)),
    inference(resolve, [], [refute_0_600, refute_0_599])).

cnf(refute_0_602, plain,
    (1 +
     X_000028 *
     (X_000028 *
      (-1/2 +
       X_000028 *
       (X_000028 *
        (1/24 +
         X_000028 *
         (X_000028 *
          (-1/720 +
           X_000028 *
           (X_000028 *
            (1/40320 + X_000028 * (X_000028 * -1/3628800))))))))) <
     X_000029 | X_000029 <= cos(X_000028)),
    inference(arithmetic, [], [refute_0_601])).

cnf(refute_0_603, plain,
    (1 +
     2443/10000 *
     (2443/10000 *
      (-1/2 +
       2443/10000 *
       (2443/10000 *
        (1/24 +
         2443/10000 *
         (2443/10000 *
          (-1/720 +
           2443/10000 *
           (2443/10000 *
            (1/40320 + 2443/10000 * (2443/10000 * -1/3628800))))))))) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     cos(2443/10000)), inference(subst, [], [refute_0_602])).

cnf(refute_0_604, plain,
    (0 < sin(2443/10000) * (3/125 + skoT_thetaC1 * 1/2) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_590])).

cnf(refute_0_605, plain,
    (0 < 0 | 3/125 + skoT_thetaC1 * 1/2 < 0 |
     0 < 3/125 + skoT_thetaC1 * 1/2 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_604])).

cnf(refute_0_606, plain,
    (skoT_thetaC1 < -6/125 | -6/125 < skoT_thetaC1 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_605])).

cnf(refute_0_607, plain,
    (-6/125 < skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_523, refute_0_606])).

cnf(refute_0_608, plain,
    (skoT_thetaC1 < 6/125 | 0 <= skoT_thetaC1 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_265, refute_0_266, refute_0_59, refute_0_50,
               refute_0_96, refute_0_157, refute_0_41, refute_0_267,
               refute_0_523, refute_0_607])).

cnf(refute_0_609, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_381, refute_0_608])).

cnf(refute_0_610, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_609])).

cnf(refute_0_611, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), inference(arithmetic, [], [refute_0_610])).

cnf(refute_0_612, plain,
    (sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <
     0 |
     -6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_609])).

cnf(refute_0_613, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) * (cos(2443/10000) * -1357143/40000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_612])).

cnf(refute_0_614, plain,
    (-1307916368539235571037117871615276716433805369757135585645641/
     192500000000000000000000000000000000000000000000000000000000 <=
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_265, refute_0_266, refute_0_59, refute_0_50,
               refute_0_96, refute_0_157, refute_0_41, refute_0_267,
               refute_0_523, refute_0_607])).

cnf(refute_0_615, plain,
    (sin(1047/5000) * (cos(2443/10000) * -1357143/40000) <=
     -6/125 + skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_614, refute_0_411])).

cnf(refute_0_616, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_615, refute_0_613])).

cnf(refute_0_617, plain,
    (6/125 <= skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_382, refute_0_418])).

cnf(refute_0_618, plain,
    (0 < 0 | -3/125 + skoT_thetaC1 * 1/2 < 0 |
     0 < -3/125 + skoT_thetaC1 * 1/2 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(split, [], [refute_0_416])).

cnf(refute_0_619, plain,
    (skoT_thetaC1 < 6/125 | 6/125 < skoT_thetaC1 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_618])).

cnf(refute_0_620, plain,
    (0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0),
    inference(decision, [],
              [refute_0_617, refute_0_619, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267, refute_0_523, refute_0_607])).

cnf(refute_0_621, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_620, refute_0_598])).

cnf(refute_0_622, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_623, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     137928037016074197461091406895432752238788286029099/
     23760000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_622])).

cnf(refute_0_624, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     skoT_thetaC1 <= -6/125 | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_623])).

cnf(refute_0_625, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 | sin(1047/5000) = 0),
    inference(resolve, [], [refute_0_624, refute_0_607])).

cnf(refute_0_626, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_625, refute_0_289])).

cnf(refute_0_627, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_621, refute_0_626])).

cnf(refute_0_628, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_31, refute_0_627])).

cnf(refute_0_629, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_628])).

cnf(refute_0_630, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_30, refute_0_629])).

cnf(refute_0_631, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_630])).

cnf(refute_0_632, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_631, refute_0_626])).

cnf(refute_0_633, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(split, [], [refute_0_632])).

cnf(refute_0_634, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_633])).

cnf(refute_0_635, plain,
    (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <
     cos(2443/10000) |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(subst, [], [refute_0_365])).

cnf(refute_0_636, plain,
    (cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(subst, [], [leq_left_mul_divide_pos])).

cnf(refute_0_637, plain,
    (1047/5000 < 0 |
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <
     sin(1047/5000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(subst, [], [refute_0_457])).

cnf(refute_0_638, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000), introduced(tautology, [assume])).

cnf(refute_0_639, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(split, [], [refute_0_638])).

cnf(refute_0_640, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     skoT_thetaC1 <= -6/125), inference(arithmetic, [], [refute_0_639])).

cnf(refute_0_641, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_640, refute_0_607])).

cnf(refute_0_642, plain,
    (1047/5000 < 0 |
     cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_641, refute_0_637])).

cnf(refute_0_643, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_642])).

cnf(refute_0_644, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(resolve, [], [refute_0_643, refute_0_636])).

cnf(refute_0_645, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     skoT_thetaC1 <= -6/125), inference(arithmetic, [], [refute_0_644])).

cnf(refute_0_646, plain,
    (-6/125 < skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0),
    inference(decision, [],
              [refute_0_617, refute_0_619, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267, refute_0_523, refute_0_607])).

cnf(refute_0_647, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_645, refute_0_646])).

cnf(refute_0_648, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0 |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_647, refute_0_635])).

cnf(refute_0_649, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0 |
     skoT_thetaC1 <=
     -3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625),
    inference(arithmetic, [], [refute_0_648])).

cnf(refute_0_650, plain,
    (-3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 <
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0),
    inference(decision, [],
              [refute_0_617, refute_0_619, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267, refute_0_523, refute_0_607])).

cnf(refute_0_651, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_649, refute_0_650])).

cnf(refute_0_652, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_634, refute_0_651])).

cnf(refute_0_653, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_29, refute_0_652])).

cnf(refute_0_654, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_653])).

cnf(refute_0_655, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_28, refute_0_654])).

cnf(refute_0_656, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_655])).

cnf(refute_0_657, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_656, refute_0_651])).

cnf(refute_0_658, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(split, [], [refute_0_657])).

cnf(refute_0_659, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_658])).

cnf(refute_0_660, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_659, refute_0_607])).

cnf(refute_0_661, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_27, refute_0_660])).

cnf(refute_0_662, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_661])).

cnf(refute_0_663, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_22, refute_0_662])).

cnf(refute_0_664, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_663])).

cnf(refute_0_665, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_664, refute_0_646])).

cnf(refute_0_666, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_21, refute_0_665])).

cnf(refute_0_667, plain,
    (5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 <
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(arithmetic, [], [refute_0_666])).

cnf(refute_0_668, plain,
    (sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0 |
     skoT_thetaC1 <=
     5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125),
    inference(decision, [],
              [refute_0_617, refute_0_619, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267, refute_0_523, refute_0_607])).

cnf(refute_0_669, plain,
    (sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_668, refute_0_667])).

cnf(refute_0_670, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_669, refute_0_616])).

cnf(refute_0_671, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * 1/2))) <=
     0), inference(resolve, [], [refute_0_670, refute_0_595])).

cnf(refute_0_672, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_671, refute_0_590])).

cnf(refute_0_673, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_672])).

cnf(refute_0_674, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_673])).

cnf(refute_0_675, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_674, refute_0_626])).

cnf(refute_0_676, plain,
    (2443/10000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_31, refute_0_675])).

cnf(refute_0_677, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_676])).

cnf(refute_0_678, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_30, refute_0_677])).

cnf(refute_0_679, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_678])).

cnf(refute_0_680, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_679, refute_0_626])).

cnf(refute_0_681, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_680])).

cnf(refute_0_682, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_681])).

cnf(refute_0_683, plain,
    (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)) <=
     0 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) <=
     0 /
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))),
    inference(split, [], [refute_0_669])).

cnf(refute_0_684, plain,
    (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)) <=
     0 | sin(1047/5000) * sin(1047/5000) <= 0 | sin(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_683])).

cnf(refute_0_685, plain,
    (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)) <=
     0 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_684, refute_0_9])).

cnf(refute_0_686, plain,
    (sin(1047/5000) * sin(1047/5000) <= 0 |
     -3/125 + skoT_thetaC1 * 1/2 <= 0 | sin(1047/5000) = 0),
    inference(split, [], [refute_0_685])).

cnf(refute_0_687, plain,
    (sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= 6/125 |
     sin(1047/5000) = 0), inference(arithmetic, [], [refute_0_686])).

cnf(refute_0_688, plain,
    (sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= 6/125),
    inference(resolve, [], [refute_0_687, refute_0_289])).

cnf(refute_0_689, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), introduced(tautology, [assume])).

cnf(refute_0_690, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) * (3/125 + skoT_thetaC1 * -1/2) <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_689])).

cnf(refute_0_691, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(2443/10000) * (3/125 + skoT_thetaC1 * -1/2) <= 0),
    inference(resolve, [], [refute_0_690, refute_0_289])).

cnf(refute_0_692, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     3/125 + skoT_thetaC1 * -1/2 <= 0 |
     sin(2443/10000) <= 0 / (3/125 + skoT_thetaC1 * -1/2)),
    inference(split, [], [refute_0_691])).

cnf(refute_0_693, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     6/125 <= skoT_thetaC1 | sin(2443/10000) <= 0),
    inference(arithmetic, [], [refute_0_692])).

cnf(refute_0_694, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     6/125 <= skoT_thetaC1),
    inference(resolve, [], [refute_0_693, refute_0_9])).

cnf(refute_0_695, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     -6/125 < skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_607, refute_0_523, refute_0_267, refute_0_41,
               refute_0_157, refute_0_96, refute_0_50, refute_0_59,
               refute_0_266, refute_0_265, refute_0_688, refute_0_670,
               refute_0_694])).

cnf(refute_0_696, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_645, refute_0_695])).

cnf(refute_0_697, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_696, refute_0_635])).

cnf(refute_0_698, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     skoT_thetaC1 <=
     -3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625),
    inference(arithmetic, [], [refute_0_697])).

cnf(refute_0_699, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     -3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 <
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_607, refute_0_523, refute_0_267, refute_0_41,
               refute_0_157, refute_0_96, refute_0_50, refute_0_59,
               refute_0_266, refute_0_265, refute_0_688, refute_0_670,
               refute_0_694])).

cnf(refute_0_700, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_698, refute_0_699])).

cnf(refute_0_701, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_682, refute_0_700])).

cnf(refute_0_702, plain,
    (1047/5000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_29, refute_0_701])).

cnf(refute_0_703, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_702])).

cnf(refute_0_704, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_28, refute_0_703])).

cnf(refute_0_705, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_704])).

cnf(refute_0_706, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_705, refute_0_700])).

cnf(refute_0_707, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_706])).

cnf(refute_0_708, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_707])).

cnf(refute_0_709, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_708, refute_0_607])).

cnf(refute_0_710, plain,
    (1047/5000 < 0 |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_27, refute_0_709])).

cnf(refute_0_711, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_710])).

cnf(refute_0_712, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_22, refute_0_711])).

cnf(refute_0_713, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_712])).

cnf(refute_0_714, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_713, refute_0_695])).

cnf(refute_0_715, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_21, refute_0_714])).

cnf(refute_0_716, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 <
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_715])).

cnf(refute_0_717, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     skoT_thetaC1 <=
     5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125),
    inference(decision, [],
              [refute_0_607, refute_0_523, refute_0_267, refute_0_41,
               refute_0_157, refute_0_96, refute_0_50, refute_0_59,
               refute_0_266, refute_0_265, refute_0_688, refute_0_670,
               refute_0_694])).

cnf(refute_0_718, plain,
    (0 <
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_717, refute_0_716])).

cnf(refute_0_719, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_611, refute_0_718])).

cnf(refute_0_720, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_332, refute_0_719])).

cnf(refute_0_721, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_720])).

cnf(refute_0_722, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(2443/10000) *
      (sin(1047/5000) *
       (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2)))) |
     0 <=
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (-3/125 + skoT_thetaC1 * 1/2))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_331, refute_0_721])).

cnf(refute_0_723, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(2443/10000) *
     (sin(1047/5000) * (sin(1047/5000) * (3/125 + skoT_thetaC1 * -1/2))) <=
     0), inference(arithmetic, [], [refute_0_722])).

cnf(refute_0_724, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     sin(2443/10000) *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_723, refute_0_718])).

cnf(refute_0_725, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_724])).

cnf(refute_0_726, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <= skoT_thetaC1 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_725])).

cnf(refute_0_727, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0),
    introduced(tautology, [assume])).

cnf(refute_0_728, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     137928037016074197461091406895432752238788286029099/
     23760000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -137928037016074197461091406895432752238788286029099/
     1140480000000000000000000000000000000000000000000000 <= 0 |
     sin(1047/5000) = 0), inference(split, [], [refute_0_727])).

cnf(refute_0_729, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     6/125 <= skoT_thetaC1 | sin(1047/5000) = 0),
    inference(arithmetic, [], [refute_0_728])).

cnf(refute_0_730, plain,
    (0 < sin(2443/10000) * (3/125 + skoT_thetaC1 * -1/2) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_718])).

cnf(refute_0_731, plain,
    (0 < 0 | 3/125 + skoT_thetaC1 * -1/2 < 0 |
     0 < 3/125 + skoT_thetaC1 * -1/2 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_730])).

cnf(refute_0_732, plain,
    (skoT_thetaC1 < 6/125 | 6/125 < skoT_thetaC1 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_731])).

cnf(refute_0_733, plain,
    (skoT_thetaC1 < 6/125 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_688, refute_0_732])).

cnf(refute_0_734, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0 | sin(1047/5000) = 0),
    inference(resolve, [], [refute_0_729, refute_0_733])).

cnf(refute_0_735, plain,
    (0 <
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_734, refute_0_289])).

cnf(refute_0_736, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_726, refute_0_735])).

cnf(refute_0_737, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_330, refute_0_736])).

cnf(refute_0_738, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_737])).

cnf(refute_0_739, plain,
    (-6/125 + skoT_thetaC1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_329, refute_0_738])).

cnf(refute_0_740, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_739])).

cnf(refute_0_741, plain,
    (-6/125 + skoT_thetaC1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_740, refute_0_735])).

cnf(refute_0_742, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_741])).

cnf(refute_0_743, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_742])).

cnf(refute_0_744, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     sin(1047/5000) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)),
    inference(resolve, [], [refute_0_461, refute_0_733])).

cnf(refute_0_745, plain,
    (1047/5000 < 0 |
     cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_744, refute_0_458])).

cnf(refute_0_746, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <=
     cos(2443/10000) * -1357143/40000 /
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_745])).

cnf(refute_0_747, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
     650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 <=
     0), inference(resolve, [], [refute_0_746, refute_0_453])).

cnf(refute_0_748, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     6/125 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(arithmetic, [], [refute_0_747])).

cnf(refute_0_749, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_748, refute_0_733])).

cnf(refute_0_750, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_749, refute_0_452])).

cnf(refute_0_751, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 <=
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_750])).

cnf(refute_0_752, plain,
    (skoT_thetaC1 <
     3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_733, refute_0_688, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267, refute_0_523, refute_0_607])).

cnf(refute_0_753, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_751, refute_0_752])).

cnf(refute_0_754, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_743, refute_0_753])).

cnf(refute_0_755, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_328, refute_0_754])).

cnf(refute_0_756, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_755])).

cnf(refute_0_757, plain,
    (-6/125 + skoT_thetaC1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_327, refute_0_756])).

cnf(refute_0_758, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_757])).

cnf(refute_0_759, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_758, refute_0_753])).

cnf(refute_0_760, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_759])).

cnf(refute_0_761, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | 6/125 <= skoT_thetaC1 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_760])).

cnf(refute_0_762, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_761, refute_0_733])).

cnf(refute_0_763, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_326, refute_0_762])).

cnf(refute_0_764, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_763])).

cnf(refute_0_765, plain,
    (-6/125 + skoT_thetaC1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_325, refute_0_764])).

cnf(refute_0_766, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     0 <= skoT_thetaC1 | 6/125 <= skoT_thetaC1 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_765])).

cnf(refute_0_767, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_766, refute_0_608])).

cnf(refute_0_768, plain,
    (1 +
     2443/10000 *
     (2443/10000 *
      (-1/2 +
       2443/10000 *
       (2443/10000 *
        (1/24 +
         2443/10000 *
         (2443/10000 *
          (-1/720 +
           2443/10000 *
           (2443/10000 *
            (1/40320 + 2443/10000 * (2443/10000 * -1/3628800))))))))) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_603, refute_0_767])).

cnf(refute_0_769, plain,
    (skoT_thetaC1 <
     -5077105619325820906765316557272753929357241166913404116960104665413505396543737332227583832964530104748665321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 |
     0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_768])).

cnf(refute_0_770, plain,
    (-5077105619325820906765316557272753929357241166913404116960104665413505396543737332227583832964530104748665321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 <=
     skoT_thetaC1 | 0 <= skoT_thetaC1 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_733, refute_0_688, refute_0_265, refute_0_266,
               refute_0_59, refute_0_50, refute_0_96, refute_0_157,
               refute_0_41, refute_0_267, refute_0_523, refute_0_607])).

cnf(refute_0_771, plain,
    (0 <= skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_770, refute_0_769])).

cnf(refute_0_772, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_771, refute_0_598])).

cnf(refute_0_773, plain,
    (sin(2443/10000) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_772, refute_0_626])).

cnf(refute_0_774, plain,
    (2443/10000 < 0 |
     2443/10000 *
     (1 +
      2443/10000 *
      (2443/10000 *
       (-1/6 +
        2443/10000 *
        (2443/10000 *
         (1/120 +
          2443/10000 *
          (2443/10000 *
           (-1/5040 +
            2443/10000 *
            (2443/10000 *
             (1/362880 + 2443/10000 * (2443/10000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_31, refute_0_773])).

cnf(refute_0_775, plain,
    (137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      sin(1047/5000) * (cos(2443/10000) * 1357143/40000)) /
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_774])).

cnf(refute_0_776, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     sin(1047/5000) * (cos(2443/10000) * 1357143/40000) <
     137928037016074197461091406895432752238788286029099/
     570240000000000000000000000000000000000000000000000 *
     (sin(1047/5000) *
      (sin(1047/5000) *
       (-137928037016074197461091406895432752238788286029099/
        23760000000000000000000000000000000000000000000000000 +
        skoT_thetaC1 *
        -137928037016074197461091406895432752238788286029099/
        1140480000000000000000000000000000000000000000000000))) |
     0 <=
     sin(1047/5000) *
     (sin(1047/5000) *
      (-137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_30, refute_0_775])).

cnf(refute_0_777, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (sin(1047/5000) *
      (137928037016074197461091406895432752238788286029099/
       23760000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       137928037016074197461091406895432752238788286029099/
       1140480000000000000000000000000000000000000000000000)) <= 0 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_776])).

cnf(refute_0_778, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     sin(1047/5000) *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_777, refute_0_626])).

cnf(refute_0_779, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_778])).

cnf(refute_0_780, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_779])).

cnf(refute_0_781, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     cos(2443/10000) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_645, refute_0_607])).

cnf(refute_0_782, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     1 + 2443/10000 * (2443/10000 * -1/2) <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     526733798302800000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     25283222318534400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000),
    inference(resolve, [], [refute_0_781, refute_0_635])).

cnf(refute_0_783, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 |
     skoT_thetaC1 <=
     -3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625),
    inference(arithmetic, [], [refute_0_782])).

cnf(refute_0_784, plain,
    (-3066119623787523080986636589038675024129329145313105239367441571373751126168960790410374650942222624504512027283312092579628830984188668733482453470382/
     566415006869060471595604972396336027690527356359153488366070286481795186683133549471894629638010510667235069001928742267312170597265280884447299625 <
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(decision, [],
              [refute_0_771, refute_0_733, refute_0_688, refute_0_265,
               refute_0_266, refute_0_59, refute_0_50, refute_0_96,
               refute_0_157, refute_0_41, refute_0_267, refute_0_523,
               refute_0_607])).

cnf(refute_0_785, plain,
    (cos(2443/10000) * -1357143/40000 <
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_783, refute_0_784])).

cnf(refute_0_786, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_780, refute_0_785])).

cnf(refute_0_787, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_29, refute_0_786])).

cnf(refute_0_788, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1) /
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_787])).

cnf(refute_0_789, plain,
    (-6/125 + skoT_thetaC1 * -1 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (cos(2443/10000) * -1357143/40000 +
      sin(1047/5000) *
      (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
       skoT_thetaC1 *
       -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
       650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000)) |
     0 <=
     cos(2443/10000) * -1357143/40000 +
     sin(1047/5000) *
     (-19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_28, refute_0_788])).

cnf(refute_0_790, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) *
     (19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      13548902400000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      19024143395107534004575823085328764995100821533462778017439060981912037912206384845610796669474751801/
      650347315200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) <=
     cos(2443/10000) * -1357143/40000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_789])).

cnf(refute_0_791, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     sin(1047/5000) *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_790, refute_0_785])).

cnf(refute_0_792, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(split, [], [refute_0_791])).

cnf(refute_0_793, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_792])).

cnf(refute_0_794, plain,
    (sin(1047/5000) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_793, refute_0_607])).

cnf(refute_0_795, plain,
    (1047/5000 < 0 |
     1047/5000 *
     (1 +
      1047/5000 *
      (1047/5000 *
       (-1/6 +
        1047/5000 *
        (1047/5000 *
         (1/120 +
          1047/5000 *
          (1047/5000 *
           (-1/5040 +
            1047/5000 *
            (1047/5000 *
             (1/362880 + 1047/5000 * (1047/5000 * -1/39916800)))))))))) <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_27, refute_0_794])).

cnf(refute_0_796, plain,
    (5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 <
     (-6/125 + skoT_thetaC1 * -1 +
      cos(2443/10000) *
      6788354801473886462206476050485555410123600697481391/
      962500000000000000000000000000000000000000000000000) /
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_795])).

cnf(refute_0_797, plain,
    (-6/125 + skoT_thetaC1 * -1 +
     cos(2443/10000) *
     6788354801473886462206476050485555410123600697481391/
     962500000000000000000000000000000000000000000000000 <
     5001945116670746164705175541918246942380869737/
     24062500000000000000000000000000000000000000000 *
     (-4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
      skoT_thetaC1 *
      -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
      745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000) |
     0 <=
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     15524784000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     -4531320054952483772764839779170688221524218850873227906928562291854361493465068395775157037104084085337880552015429938138497364778122247075578397/
     745189632000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_22, refute_0_796])).

cnf(refute_0_798, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0 | skoT_thetaC1 <= -6/125),
    inference(arithmetic, [], [refute_0_797])).

cnf(refute_0_799, plain,
    (cos(2443/10000) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_798, refute_0_607])).

cnf(refute_0_800, plain,
    (1 + 2443/10000 * (2443/10000 * -1/2) <
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     878231183402041408052330867374677857185002617755399494121200000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 +
     skoT_thetaC1 *
     5969486701926352735544279544027281418434957487134346215885378416232243463437896301570718717266697285639113134224507007728156425689954426515535985044757597940474006401743352477825321154270576137/
     42155096803297987586511881633984537144880125652259175717817600000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000000 |
     sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_21, refute_0_799])).

cnf(refute_0_801, plain,
    (5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125 <
     skoT_thetaC1 | sin(1047/5000) * sin(1047/5000) <= 0),
    inference(arithmetic, [], [refute_0_800])).

cnf(refute_0_802, plain,
    (sin(1047/5000) * sin(1047/5000) <= 0 |
     skoT_thetaC1 <=
     5076325108737448824953712057577561842205409920957109529170206996138606539219372622190575687696399816286165321194652957953631061445860273440906784089731454412357155961589539885133048073074376543178/
     746185837740794091943034943003410177304369685891793276985672302029030432929737037696339839658337160704889141778063375966019553211244303314441998130594699742559250800217919059728165144283822017125),
    inference(decision, [],
              [refute_0_771, refute_0_733, refute_0_688, refute_0_265,
               refute_0_266, refute_0_59, refute_0_50, refute_0_96,
               refute_0_157, refute_0_41, refute_0_267, refute_0_523,
               refute_0_607])).

cnf(refute_0_803, plain, (sin(1047/5000) * sin(1047/5000) <= 0),
    inference(resolve, [], [refute_0_802, refute_0_801])).

cnf(refute_0_804, plain, (0 < sin(1047/5000) * (sin(1047/5000) * -1)),
    inference(resolve, [], [refute_0_803, refute_0_16])).

cnf(refute_0_805, plain, (0 < -1), inference(split, [], [refute_0_804])).

cnf(refute_0_806, plain, ($false),
    inference(arithmetic, [], [refute_0_805])).
SZS output end CNFRefutation for /home/omar/Desktop/Omar - ex/My Problems/Helicopter/Torque_theta.tptp

Processor time: 24.908 = 15.432 (Metis) + 9.476 (RCF)
Maximum weight in proof search: 542



